```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08A.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session #08-A <br></center></h2>

<h1><center>Numerical differentiation<br></center></h1>

---

```python
from resources.hashutils import *
import numpy as np
import matplotlib.pyplot as plt
```

# Question 1: Ant on a wire

A camera records the position of an ant as it walks on a wire. The ant generally moves forward (direction of increasing $x$ in the figure), but sometimes briefly turns around when it bumps into an oncoming ant. A researcher has processed the video and extracted a sequence of positions with their respective time stamps. In this exercise we will analyze this data to estimate the speed of the ant and count the number of times that it moves backwards. 


<img src="resources/ants.png" width="900"/>

## Question 1.1: Load the data

The data is stored in two files called `times.npy` and `positions.npy`. Load these two files into NumPy arrays called `time` and `xpos` respectively. 

```python
time = ...
xpos = ...
```

```python
grader.check("q1p1")
```

## Question 1.2: Plot the data

Plot the measured positions versus time. The plot should have these characteristics:

+ `figsize=(10,4)`
+ Use the `ax1.plot` function.
+ Use a `.` marker for the plot.
+ Add a grid.
+ Label the x-axis with `'time [sec]'`
+ Label the y-axis with `'position [mm]'`
+ Both x-axis and y-axis labels should have a font size of 16.
+ Remove the top and right-hand spines with  
```python 
ax1.spines[['top','right']].set_visible(False)
```

```python
fig1, ax1 = plt.subplots(...)
...
...
...
...
```

```python
grader.check("q1p2")
```

## Question 1.3: Estimate the time step $h$

The time step $h$ is the interval of time between time stamps in `time` array. These time stamps are uniformly spaced. 

```python
h = ...
```

```python
grader.check("q1p3")
```

## Question 1.4 Forward differencing

Write a function `diff_forward(f,h)` that computes the forward difference estimate of the derivative of a function $f(t)$ from uniformly spaced evaluations of the function. 

$$f'_F(t) = \frac{f(t+h)-f(t)}{h}$$

The function `diff_forward(f,h)` takes as inputs:
+ `f`: a 1D NumPy array of function evaluations, and 
+ `h`: the step size. 

The function should return a NumPy array whose shape is the same as that of `f`, and whose last entry is `np.nan`. 

```python
def diff_forward(f,h):
    ...
    ...
    ...
    return ...
```

```python
grader.check("q1p4")
```

## Question 1.5: Backward differencing

Write a function `diff_backward(f,h)` that computes the backward difference estimate of the derivative of a function $f(t)$ from uniformly spaced evaluations of the function. 

$$f'_B(t) = \frac{f(t)-f(t-h)}{h}$$

The input to this function is the same as to `diff_backward(f,h)`. The output is also similar, except that here the first (as opposed to the last) entry of the output array should be `np.nan`. 

```python
def diff_backward(f,h):
    ...
    ...
    ...
    return ...
```

```python
grader.check("q1p5")
```

## Question 1.6: Central differencing

Write a function `diff_central(f,h)` that computes the central difference estimate of the derivative of a function $f(t)$ from uniformly spaced evaluations of the function. 

$$f'_C(t) = \frac{f(t+h)-f(t-h)}{2h}$$

The input to this function is the same as the previous two questions. The output is an array of the same shape as `f`, whose first and last entries are `np.nan`. 

```python
def diff_central(f,h):
    ...
    ...
    ...
    ...
    return ...
```

```python
grader.check("q1p6")
```

## Question 1.7:  Plot the three speed estimates

Plot the three estimates of the ant's speed (forward, backward, and central differencing) versus time on a single plot. 
The characteristics of the plot are as follows:

+ `figsize=(8,4)`
+ Use a boolean mask to select the portions of `vB`, `vF`, and `vC` between time `time>=10` and `time<=20`, and plot only those portions of the velocity estimates. 
+ Label the x-axis with `'time [sec]'`
+ Label the y-axis with `'speed [mm/sec]'`
+ Both x-axis and y-axis labels should have a font size of 16.
+ Label the plots as `'Forward'`, `'Backward'`, or `'Central'`, and add a legend. 
+ Add grid lines to the plot. 

```python
# ****** Don't modify this cell *******
vF = diff_forward(xpos,h)
vB = diff_backward(xpos,h)
vC = diff_central(xpos,h)
# *************************************
```

```python
fig2, ax2 = plt.subplots(...)
...
...  # etc.
...
```

```python
grader.check("q1p7")
```

## Question 1.8: Count the number of times that the ant goes from moving forward to moving backward. 

Write a function called `count_turn_back` that takes the estimated speed profile as input and returns the number of times that it goes from a positive value to a negative value. 

```python
def count_turn_backs(v):
    ...
    return ...
```

```python
grader.check("q1p8")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1p1 (points: None)
```
>>> time.shape == (200,) and xpos.shape == (200,)
True
```
```
>>> get_hash(xpos, 5) == 'd9bfc888ddb40e739b397985b8de72c4' and get_hash(time, 3) == 'c954b33fbb825fd2f12fba9ec9174092'
True
```


Tests for q1p2 (points: None)
```
>>> import matplotlib
>>> isinstance(ax1, matplotlib.axes._axes.Axes)
True
```
```
>>> len(ax1.get_lines()) == 1
True
```
```
>>> line = ax1.get_lines()[0]
>>> line.get_marker() == '.' and np.all(np.isclose(line.get_data()[0], time, atol=0.01)) and np.all(np.isclose(line.get_data()[1], xpos, atol=0.01))
True
```
```
>>> line = ax1.get_lines()[0]
>>> xdata = line.get_data()[0]
>>> ydata = line.get_data()[1]
>>> np.all(np.isclose(xdata, time, atol=0.01)) and np.all(np.isclose(ydata, xpos, atol=0.01))
True
```
```
>>> ax1.get_xlabel() == 'time [sec]' and ax1.get_ylabel() == 'position [mm]'
True
```
```
>>> np.all(np.isclose(fig1.get_size_inches(), (10, 4), atol=0.01))
True
```
```
>>> ax1.get_ygridlines()[1].get_visible() and ax1.get_xgridlines()[1].get_visible()
True
```
```
>>> not ax1.spines['top'].get_visible() and (not ax1.spines['right'].get_visible())
True
```


Tests for q1p3 (points: None)
```
>>> isinstance(h, float) and h > 0 and (h < 3)
True
```
```
>>> get_hash(h, 3) == 'a7473b0d7fc57e553ee66811867c7c5a'
True
```


Tests for q1p4 (points: None)
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_forward(f, 2)
>>> fprime.shape == (10,)
True
```
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_forward(f, 2)
>>> np.all(np.isclose(fprime[:-1], 1 / 9, atol=0.001)) and np.isnan(fprime[-1])
True
```
```
>>> vF = diff_forward(xpos, h)
>>> get_hash(vF, 4) == '8041c66b986bd4671a64e3b15fab9e8b'
True
```


Tests for q1p5 (points: None)
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_backward(f, 2)
>>> fprime.shape == (10,)
True
```
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_backward(f, 2)
>>> np.all(np.isclose(fprime[1:], 1 / 9, atol=0.001)) and np.isnan(fprime[0])
True
```
```
>>> vB = diff_backward(xpos, h)
>>> get_hash(vB, 4) == 'e1d589bae5aab6f68a35ebb14bdcbc01'
True
```


Tests for q1p6 (points: None)
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_central(f, 2)
>>> fprime.shape == (10,)
True
```
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_central(f, 2)
>>> np.all(np.isclose(fprime[1:-1], 1 / 9, atol=0.001)) and np.isnan(fprime[0]) and np.isnan(fprime[0])
True
```
```
>>> vC = diff_central(xpos, h)
>>> get_hash(vC, 4) == '1bc29467455489116427c47d520222d9'
True
```


Tests for q1p7 (points: None)
```
>>> import matplotlib
>>> isinstance(ax2, matplotlib.axes._axes.Axes)
True
```
```
>>> len(ax2.get_lines()) == 3
True
```
```
>>> np.all([np.all(np.isclose(line.get_xdata(), time[(time >= 10) & (time <= 20)], atol=0.01)) for line in ax2.get_lines()])
True
```
```
>>> {line.get_label() for line in ax2.get_lines()} == {'Backward', 'Central', 'Forward'}
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> v = diff_backward(xpos, h)
>>> ind = (time >= 10) & (time <= 20)
>>> np.all(np.isclose(plotdata['Backward'], v[ind], atol=0.01))
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> v = diff_forward(xpos, h)
>>> ind = (time >= 10) & (time <= 20)
>>> np.all(np.isclose(plotdata['Forward'], v[ind], atol=0.01))
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> v = diff_central(xpos, h)
>>> ind = (time >= 10) & (time <= 20)
>>> np.all(np.isclose(plotdata['Central'], v[ind], atol=0.01))
True
```
```
>>> ax2.get_xlabel() == 'time [sec]' and ax2.get_ylabel() == 'speed [mm/sec]'
True
```
```
>>> np.all(np.isclose(fig2.get_size_inches(), (8, 4), atol=0.01))
True
```
```
>>> ax2.get_ygridlines()[1].get_visible() and ax2.get_xgridlines()[1].get_visible()
True
```
```
>>> leg = ax2.get_legend()
>>> leg.get_visible() and {t.get_text() for t in leg.get_texts()} == {'Backward', 'Central', 'Forward'}
True
```


Tests for q1p8 (points: None)
```
>>> count_turn_backs(np.array([1, 1, -1, -0.5, 1])) == 1
True
```
```
>>> count_turn_backs(vF) == count_turn_backs(vB) == count_turn_backs(vC)
True
```
```
>>> count_turn_backs(vF) >= 0 and count_turn_backs(vF) < 20
True
```
```
>>> get_hash(count_turn_backs(np.array([1, -1, 1, -1, 1, -1, 1, -1, 1])), 2) == '07078a97d66756f213dbca3e379bf084'
True
```
