```python
# Initialize Otter
import otter
grader = otter.Notebook("w07Alab.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session: Week 7-A <br></center></h2>

<h1><center>Matplotlib<br></center></h1>

---

In this lab we will practice using Matplotlib to create two-dimensional plots, including scatter plots, line plots, histograms, and bar plots. In each case you will be given a detailed specification for the plot, as well as a picture of what your plot should look like. The autograder will check that the specification has been met. 

```python
from resources.hashutils import *
import numpy as np
from pathlib import Path
import matplotlib.pyplot as plt
```

# Part 1: Scatter plot

We begin by loading data from a CSV file and visualizing it as a scatter plot. The data is a fictitious set of measurements of the water level in a tank through time. 

## Part 1.1: Import the data

The data is contained in the file `2Ddata.csv` in the `resources` folder. Use [`np.loadtxt`](https://numpy.org/doc/stable/reference/generated/numpy.loadtxt.html) to read this data into a NumPy array called `data`.  The first column of the data set is time in hours, and will be plotted along the horizontal (x) axis. The second column is the water level in meters, and will be plotted along the vertical (y) axis.

**Note**: The data in the file is comma-separated. You will have to use the `delimiter` argument of the `loadtxt` method to specify this.

```python
data = ...
```

```python
grader.check("p1p1")
```

## Part 1.2: Plot

Next, create a scatter plot of the data with the following characteristics.

+ Set the figure size to `(10,5)`.
+ Use the [`ax.scatter`](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.scatter.html) function to plot the data. 
+ Set the marker color to `'dodgerblue'`.
+ Set the y-axis limits to `(-2,2)`.
+ Set the ticks for the y axis to `[-2,-1,0,1,2]`.
+ Set the font size for both the x and y tick labels to 14. 
+ Add a grid with dashed lines (`'--'`)
+ Set the x axis label to `'time [minutes]'`.
+ Set the y axis label to `'height [m]'`.
+ Both labels should have font size equal to 18.
+ Set the visibility of the top and right spines to `False`.
+ Set the title to `'Measured water level'` with a font size of 20.

**Desired result:**

<img src="resources/scatter.png" width="600" />

```python
fig1, ax1 = plt.subplots(...)
...
```

```python
grader.check("p1p2")
```

# Part 2: Line plot

Next we will add two smooth polynomial approximations to the data from the previous question. One will be a linear function, the other a cubic function. The formulas for each of these are:

\begin{align*}
&y1 = -1.6 + 0.4\:x & \text{... linear approximation} \\
&y3 = 3.65 -4.78\: x + 1.46\: x^2 -0.12\: x^3  & \text{... cubic approximation}
\end{align*}



## Part 2.1: Compute the two approximations

First we must compute the values of y1 and y3 on an array of points along the x axis. This array should have 20 equally spaced points, and it should span from x=1 to x=7. Call the array `xx`. Then evaluate the linear and cubic functions on `xx` and call the resulting arrays `y1` and `y3` respectively. Each of these should have the same shape as `xx`.

**Hint**: [`np.linspace`](https://numpy.org/doc/stable/reference/generated/numpy.linspace.html)

```python
xx = ...
y1 = ...
y3 = ...
```

```python
grader.check("p2p1")
```

## Part 2.2: Plot

Now we can build the plot by overlaying the two lines on the scatter plot. Here is the specification:

+ All of the specifications for the scatter plot in Question 1 apply here. 

+ Label the scatter plotted data as 'data' in the legend

+ Use the [ax.plot](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.plot.html) function to plot the two lines.
+ The linear approximation should have these characteristics:
    + a line width of 2
    + dashed  line style(`'--'`)
    + star marker (`'*'`) of size 10
    + `magenta` in color
    + labels as 'linear' in the legend

+ The cubic approximation should have these characteristics:
    + a line width of 2
    + dash-dotted line style(`'-.'`)
    + round marker (`'o'`) of size 10, with a black edge
    + `orangered` in color
    + labels as 'cubic' in the legend

**Desired result:**

<img src="resources/scatter_and_line.png" width="800" />

```python
fig2, ax2 = plt.subplots(...)
...
```

```python
grader.check("p2p2")
```

# Part 3: Histogram

The data in this problem consists of the 200 scores by two teams in a popular online game. The scores are integers ranging from 0 to 100. They are stored in two `.npy` files called `team1.npy` and `team2.npy`. Your task is to present the scores as two histograms -- one for Team 1 and another for Team 2. 

## Part 3.1: Import the data.

Import the files `'team1.npy'` and `'team2.npy'` to NumPy array called `team1` and `team2` respectively.

**Hint**: For binary NumPy files we use [`np.load`](https://numpy.org/doc/stable/reference/generated/numpy.load.html)

```python
team1 = ...
team2 = ...
```

```python
grader.check("p3p1")
```

## Part 3.2: Plot

The two histograms should be placed in a figure in two rows, with Team 1 on the top row and Team 2 on the bottom row. Here are the specifications for the figure. 

+ Set the figure size to `(6,8)` and the number of rows to 2.
+ Use [`ax.hist`](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.hist.html) for both plots.
+ Top row: Team 1
    + The histogram should have 15 bins, the bar widths should be 80% of the bin size, and the bar color should be 'darkgreen'.
    + Remove the left, top, and right side spines.
    + The x-axis label should be 'Team 1' and have a font size of 16.
+ Bottom row: Team 2
    + The histogram should have 10 bins, the bar widths should be 60% of the bin size, and the bar color should be 'darkred'.
    + Remove the left, top, and right side spines.
    + The x-axis label should be 'Team 2' and have a font size of 16.

**Desired result:**

<img src="resources/histogram.png" width="400" />

```python
fig3, ax3 = plt.subplots(...)
...
```

```python
grader.check("p3p2")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for p1p1 (points: None)
```
>>> data.shape == (100, 2)
True
```
```
>>> get_hash(data, 4) == 'f9ee89b8e7b878d2e6f84ed79654d13c'
True
```


Tests for p1p2 (points: None)
```
>>> import matplotlib
>>> isinstance(ax1, matplotlib.axes._axes.Axes)
True
```
```
>>> np.all(np.isclose(ax1.collections[0].get_offsets()[:, 0], data[:, 0], atol=0.1)) and np.all(np.isclose(ax1.collections[0].get_offsets()[:, 1], data[:, 1], atol=0.1))
np.True_
```
```
>>> np.all(np.isclose(fig1.get_size_inches(), (10, 5)))
np.True_
```
```
>>> np.all(np.isclose(ax1.collections[0].get_facecolors(), np.array([[0.11764706, 0.56470588, 1.0, 1.0]]), atol=0.1))
np.True_
```
```
>>> np.all(np.isclose(ax1.get_ylim(), np.array([-2.0, 2.0]), atol=0.1))
np.True_
```
```
>>> np.all(np.isclose(ax1.get_yticks(), np.array([-2, -1, 0, 1, 2]), atol=0.1))
np.True_
```
```
>>> np.all([x.get_fontsize() == 14 for x in ax1.get_xticklabels()]) and np.all([x.get_fontsize() == 14 for x in ax1.get_yticklabels()])
np.True_
```
```
>>> ax1.get_xgridlines()[0].get_linestyle() == '--'
True
```
```
>>> ax1.get_xlabel().strip() == 'time [minutes]'
True
```
```
>>> ax1.get_ylabel().strip() == 'height [m]'
True
```
```
>>> ax1.xaxis.get_label().get_fontsize() == 18 and ax1.yaxis.get_label().get_fontsize() == 18
True
```
```
>>> np.all([not ax1.spines[x].get_visible() for x in ['top', 'right']])
np.True_
```
```
>>> ax1.get_title() == 'Measured water level' and ax1.title.get_fontsize() == 20
True
```


Tests for p2p1 (points: None)
```
>>> xx.shape == (20,) and np.all(xx >= 1) and np.all(xx <= 7)
np.True_
```
```
>>> y1.shape == (20,) and np.min(y1) < -1 and (np.max(y1) > 1)
np.True_
```
```
>>> y3.shape == (20,) and np.min(y3) < -1 and (np.max(y3) > 1.5)
np.True_
```
```
>>> np.all(np.isclose(np.diff(xx), 0.31578947)) and np.all(np.isclose(np.diff(y1), 0.12631579, atol=0.1))
np.True_
```
```
>>> np.isclose(np.sum(y3), 5.526315789473633) and np.isclose(np.mean(np.log(y3[y3 > 0])), -0.3891635509624931, atol=0.01)
np.True_
```


Tests for p2p2 (points: None)
```
>>> import matplotlib
>>> isinstance(ax2, matplotlib.axes._axes.Axes)
True
```
```
>>> len(ax2.get_lines()) == 2 and {line.get_label() for line in ax2.get_lines()} == {'linear', 'cubic'} and (ax2.collections[0].get_label() == 'data')
True
```
```
>>> line1 = list(filter(lambda x: x.get_label() == 'linear', ax2.get_lines()))[0]
>>> line1.get_linewidth() == 2
True
```
```
>>> line1 = list(filter(lambda x: x.get_label() == 'linear', ax2.get_lines()))[0]
>>> line1.get_linestyle() == '--'
True
```
```
>>> line1 = list(filter(lambda x: x.get_label() == 'linear', ax2.get_lines()))[0]
>>> line1.get_marker() == '*'
True
```
```
>>> line1 = list(filter(lambda x: x.get_label() == 'linear', ax2.get_lines()))[0]
>>> line1.get_color() == 'magenta'
True
```
```
>>> line3 = list(filter(lambda x: x.get_label() == 'cubic', ax2.get_lines()))[0]
>>> line3.get_linewidth() == 2
True
```
```
>>> line3 = list(filter(lambda x: x.get_label() == 'cubic', ax2.get_lines()))[0]
>>> line3.get_linestyle() == '-.'
True
```
```
>>> line3 = list(filter(lambda x: x.get_label() == 'cubic', ax2.get_lines()))[0]
>>> line3.get_marker() == 'o'
True
```
```
>>> line3 = list(filter(lambda x: x.get_label() == 'cubic', ax2.get_lines()))[0]
>>> line3.get_color() == 'orangered'
True
```


Tests for p3p1 (points: None)
```
>>> team1.shape == (200,) and team2.shape == (200,)
True
```
```
>>> np.isclose(np.sum(team1), 11970.743322957112, atol=0.01) and np.isclose(team1.mean(), 59.85371661478556, atol=0.01)
np.True_
```
```
>>> np.isclose(np.sum(team2), 14404.846860229123, atol=0.01) and np.isclose(team2.mean(), 72.02423430114561, atol=0.01)
np.True_
```


Tests for p3p2 (points: None)
```
>>> np.all(np.isclose(fig3.get_size_inches(), (6, 8)))
np.True_
```
```
>>> len(ax3) and ax3[0].get_gridspec().get_geometry() == (2, 1)
True
```
```
>>> import matplotlib
>>> sum([np.all(np.isclose(c.get_facecolor(), (0.0, 0.3921, 0.0, 1.0), atol=0.01)) for c in ax3[0].get_children() if isinstance(c, matplotlib.patches.Rectangle)]) == 15
np.True_
```
```
>>> import matplotlib
>>> sum([np.isclose(c.get_width(), 5.2384, atol=0.05) for c in ax3[0].get_children() if isinstance(c, matplotlib.patches.Rectangle)]) == 15
np.True_
```
```
>>> np.all([not ax3[0].spines[x].get_visible() for x in ['left', 'top', 'right']])
np.True_
```
```
>>> ax3[0].get_xlabel().strip() == 'Team 1' and ax3[0].xaxis.get_label().get_fontsize() == 16
True
```
```
>>> import matplotlib
>>> sum([np.all(np.isclose(c.get_facecolor(), (0.54509, 0.0, 0.0, 1.0), atol=0.01)) for c in ax3[1].get_children() if isinstance(c, matplotlib.patches.Rectangle)]) == 10
np.True_
```
```
>>> import matplotlib
>>> sum([np.isclose(c.get_width(), 3.4121514, atol=0.05) for c in ax3[1].get_children() if isinstance(c, matplotlib.patches.Rectangle)]) == 10
np.True_
```
```
>>> np.all([not ax3[1].spines[x].get_visible() for x in ['left', 'top', 'right']])
np.True_
```
```
>>> ax3[1].get_xlabel().strip() == 'Team 2' and ax3[1].xaxis.get_label().get_fontsize() == 16
True
```
