```python
# Initialize Otter
import otter
grader = otter.Notebook("w11Alab.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session: Week 11-A <br></center></h2>


<h1><center>Numerical differentiation<br></center></h1>

---

```python
from resources.hashutils import *
from pathlib import Path
import numpy as np
import matplotlib.pyplot as plt
```

# Part 1: Ant on a wire

A camera records the position of an ant as it walks on a wire. The ant generally moves forward (direction of increasing $x$ in the figure), but sometimes briefly turns around when it bumps into an oncoming ant. A researcher has processed the video and extracted a sequence of positions with their respective time stamps. In this exercise we will analyze this data to estimate the speed of the ant and count the number of times that it moves backwards. 

<center><img src="resources/ants.png" width="600"/></center>

## Part 1.1: Load the data

The data is stored in two files called `times.npy` and `positions.npy`. Both are located in the `resources` foleder. Use [np.load](https://numpy.org/doc/stable/reference/generated/numpy.load.html) to load these two files into NumPy arrays called `time` and `xpos` respectively. 



```python
...
```

```python
grader.check("p1p1")
```

## Part 1.2: Plot the data

Plot the measured positions versus time. The plot should have these characteristics:

+ Use variable `fig1` for the figure and `ax1` for the axes.
+ Set the figure size to `(8,3)`.
+ Use the `plot` function to create the plot. 
+ Set the width of the line to `3`.
+ Add a grid.
+ Label the x-axis with `'time [sec]'`
+ Label the y-axis with `'position [mm]'`
+ Both x-axis and y-axis labels should have a font size of 14.
+ Remove the top and right-hand spines.

```python
...
```

```python
grader.check("p1p2")
```

## Part 1.3: Compute the time step $h$

The time step $h$ is the interval of time between time stamps in `time` array. These time stamps are uniformly spaced. Find the value of $h$ and set it to the variable name `h`.

```python
...
```

```python
grader.check("p1p3")
```

## Part 1.4 Forward differencing

Write a function `diff_forward(f,h)` that computes the forward difference estimate of the derivative of a function $f(t)$ from uniformly spaced evaluations of the function. 

$$f'_F(t) = \frac{f(t+h)-f(t)}{h}$$

The function `diff_forward(f,h)` takes as inputs:
+ `f`: a 1D NumPy array of uniformly spaced function evaluations, and 
+ `h`: the step size.

The function should return a NumPy array whose shape is the same as that of `f`, and whose last entry is `np.nan`. 

```python
...
```

```python
grader.check("p1p4")
```

## Part 1.5: Backward differencing

Write a function `diff_backward(f,h)` that computes the backward difference estimate of the derivative of a function $f(t)$ from uniformly spaced evaluations of the function. 

$$f'_B(t) = \frac{f(t)-f(t-h)}{h}$$

The input to this function is the same as to `diff_backward(f,h)`. The output is also similar, except that here the first (as opposed to the last) entry of the output array should be `np.nan`. 

```python
...
```

```python
grader.check("p1p5")
```

## Part 1.6: Central differencing

Write a function `diff_central(f,h)` that computes the central difference estimate of the derivative of a function $f(t)$ from uniformly spaced evaluations of the function. 

$$f'_C(t) = \frac{f(t+h)-f(t-h)}{2h}$$

The input to this function is the same as the previous two questions. The output is an array of the same shape as `f`, whose first and last entries are `np.nan`. 

```python
...
```

```python
grader.check("p1p6")
```

## Part 1.7:  Plot the three speed estimates

Plot the three estimates of the ant's speed (forward, backward, and central differencing) versus time on a single plot. 
The characteristics of the plot are as follows:

+ Use variable `fig2` for the figure and `ax2` for the axes.
+ Set the figure size to `(8,3)`.

+ Use a boolean mask to select the portions of `vF`, `vB`, and `vC` between time `time>=10` and `time<=20`, and plot only those portions of the velocity estimates. 
+ Label the x-axis with `'time [sec]'`
+ Label the y-axis with `'speed [mm/sec]'`
+ Both x-axis and y-axis labels should have a font size of 14.
+ Label the plots as `'Forward'`, `'Backward'`, or `'Central'`, and add a legend. 
+ Add grid lines to the plot. 

```python
# ****** Don't modify this cell *******
vF = diff_forward(xpos,h)
vB = diff_backward(xpos,h)
vC = diff_central(xpos,h)
# *************************************
```

```python
...
```

```python
grader.check("p1p7")
```

## Part 1.8: Count the number of times that the ant goes from moving forward to moving backward. 

Write a function called `count_turn_backs` that takes a speed profile `v` as input and returns the number of times that it goes from a positive value to a negative value. Apply your function to the three speed estimates. Is the number of turns consistent among all three speed estimators?

```python
def count_turn_backs(v):
    ...
    return ...
```

```python
grader.check("p1p8")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for p1p1 (points: None)
```
>>> 'time' in globals()
True
```
```
>>> 'xpos' in globals()
True
```
```
>>> time.shape == (200,)
True
```
```
>>> xpos.shape == (200,)
True
```
```
>>> get_hash(time, 3) == 'c954b33fbb825fd2f12fba9ec9174092'
True
```
```
>>> get_hash(xpos, 5) == 'd9bfc888ddb40e739b397985b8de72c4'
True
```


Tests for p1p2 (points: None)
```
>>> 'ax1' in globals() and 'fig1' in globals()
True
```
```
>>> import matplotlib
>>> isinstance(ax1, matplotlib.axes._axes.Axes)
True
```
```
>>> import matplotlib
>>> isinstance(fig1, matplotlib.figure.Figure)
True
```
```
>>> np.all(np.isclose(fig1.get_size_inches(), (8, 3), atol=0.01))
np.True_
```
```
>>> len(ax1.get_lines()) == 1
True
```
```
>>> line = ax1.get_lines()[0]
>>> xdata = line.get_data()[0]
>>> ydata = line.get_data()[1]
>>> np.all(np.isclose(xdata, time, atol=0.01)) and np.all(np.isclose(ydata, xpos, atol=0.01))
np.True_
```
```
>>> ax1.get_lines()[0].get_lw() == 3
True
```
```
>>> ax1.get_xlabel() == 'time [sec]' and ax1.get_ylabel() == 'position [mm]'
True
```
```
>>> ax1.xaxis.get_label().get_fontsize() == 14 and ax1.yaxis.get_label().get_fontsize() == 14
True
```
```
>>> ax1.get_ygridlines()[1].get_visible() and ax1.get_xgridlines()[1].get_visible()
True
```
```
>>> not ax1.spines['top'].get_visible() and (not ax1.spines['right'].get_visible())
True
```


Tests for p1p3 (points: None)
```
>>> isinstance(h, float) and h > 0 and (h < 3)
np.True_
```
```
>>> get_hash(h, 3) == 'a7473b0d7fc57e553ee66811867c7c5a'
True
```


Tests for p1p4 (points: None)
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_forward(f, 2)
>>> fprime.shape == (10,)
True
```
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_forward(f, 2)
>>> np.all(np.isclose(fprime[:-1], 1 / 9, atol=0.001)) and np.isnan(fprime[-1])
np.True_
```
```
>>> vF = diff_forward(xpos, h)
>>> get_hash(vF, 4) == '8041c66b986bd4671a64e3b15fab9e8b'
True
```


Tests for p1p5 (points: None)
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_backward(f, 2)
>>> fprime.shape == (10,)
True
```
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_backward(f, 2)
>>> np.all(np.isclose(fprime[1:], 1 / 9, atol=0.001)) and np.isnan(fprime[0])
np.True_
```
```
>>> vB = diff_backward(xpos, h)
>>> get_hash(vB, 4) == 'e1d589bae5aab6f68a35ebb14bdcbc01'
True
```


Tests for p1p6 (points: None)
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_central(f, 2)
>>> fprime.shape == (10,)
True
```
```
>>> f = np.linspace(0, 2, 10)
>>> fprime = diff_central(f, 2)
>>> np.all(np.isclose(fprime[1:-1], 1 / 9, atol=0.001)) and np.isnan(fprime[0]) and np.isnan(fprime[0])
np.True_
```
```
>>> vC = diff_central(xpos, h)
>>> get_hash(vC, 4) == '1bc29467455489116427c47d520222d9'
True
```


Tests for p1p7 (points: None)
```
>>> import matplotlib
>>> isinstance(ax2, matplotlib.axes._axes.Axes)
True
```
```
>>> len(ax2.get_lines()) == 3
True
```
```
>>> np.all([np.all(np.isclose(line.get_xdata(), time[(time >= 10) & (time <= 20)], atol=0.01)) for line in ax2.get_lines()])
np.True_
```
```
>>> {line.get_label() for line in ax2.get_lines()} == {'Backward', 'Central', 'Forward'}
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> v = diff_backward(xpos, h)
>>> ind = (time >= 10) & (time <= 20)
>>> np.all(np.isclose(plotdata['Backward'], v[ind], atol=0.01))
np.True_
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> v = diff_forward(xpos, h)
>>> ind = (time >= 10) & (time <= 20)
>>> np.all(np.isclose(plotdata['Forward'], v[ind], atol=0.01))
np.True_
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> v = diff_central(xpos, h)
>>> ind = (time >= 10) & (time <= 20)
>>> np.all(np.isclose(plotdata['Central'], v[ind], atol=0.01))
np.True_
```
```
>>> ax2.get_xlabel() == 'time [sec]' and ax2.get_ylabel() == 'speed [mm/sec]'
True
```
```
>>> np.all(np.isclose(fig2.get_size_inches(), (8, 3), atol=0.01))
np.True_
```
```
>>> ax2.get_ygridlines()[1].get_visible() and ax2.get_xgridlines()[1].get_visible()
True
```
```
>>> leg = ax2.get_legend()
>>> leg.get_visible() and {t.get_text() for t in leg.get_texts()} == {'Backward', 'Central', 'Forward'}
True
```
```
>>> ax2.xaxis.get_label().get_fontsize() == 14 and ax2.yaxis.get_label().get_fontsize() == 14
True
```


Tests for p1p8 (points: None)
```
>>> count_turn_backs(np.array([1, 1, -1, -0.5, 1])) == 1
np.True_
```
```
>>> count_turn_backs(vF) == count_turn_backs(vB) == count_turn_backs(vC)
np.True_
```
```
>>> count_turn_backs(vF) >= 0 and count_turn_backs(vF) < 20
np.True_
```
```
>>> get_hash(count_turn_backs(np.array([1, -1, 1, -1, 1, -1, 1, -1, 1])), 2) == '07078a97d66756f213dbca3e379bf084'
True
```
