```python
# Initialize Otter
import otter
grader = otter.Notebook("w03Blab.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session: 03-B <br></center></h2>

<h1><center>Iterating over collections with for loops<br></center></h1>

---

```python
from resources.hashutils import *
import csv
```

# Question 1: Counting votes (contd.)

In lab 03A we found that the number of unique names of voters was less than the total number of votes collected. If we assume that no two voters have the same name (which we will), this means that at least one voter must have voted twice (by the [pigeonhole principle](https://en.wikipedia.org/wiki/Pigeonhole_principle)).We will use a `for` loop to find the offender(s).

The next cell loads the data.

```python
with open('resources/voters.csv', 'r') as file:
   voters = list(csv.reader(file))
voters = voters[0]     
with open('resources/preferred_candidates.csv', 'r') as file:
   preferred_candidates = list(csv.reader(file))
preferred_candidates = preferred_candidates[0] 
```

## Question 1.1

The first task is to count the number of votes cast by each individual voter. This data is best kept in a dictionary that maps each voter's name (the key, a string) to their number of votes (the value, an integer).

Complete the code in the cell below. 
1. Set `unique_voters` to the set of unique voter names. 
2. Initialize the dictionary `voter2numvotes` using the [`dict.fromkeys(k,value)`](https://docs.python.org/2/library/stdtypes.html#dict.fromkeys) method. This method creates a dictionary with keys taken from the set `k` and sets all values to the scalar `value`. In our case we want to pass `unique_voters` in the `k` position, and `0` for `value`.
3. Write a `for` loop that counts the number of times that each name appears in the list of `voters`. Once the loop completes, each key in `voter2numvotes` should map to the number of votes for that name. 

For example, you should have:

`voter2numvotes['Melotti Devletsah'] == 1`

```python
unique_voters = ...
voter2numvotes = ...
# Iterate through voters
for ... 	
	# Increase the count for that voter by 1.
	voter2numvotes[...] ...
```

```python
grader.check("q1p1")
```

## Question 1.2

Create a set called `multivote_voters` that collects the names of all voters who voted more than once. For this, you should iterate over the key-value pairs in `voter2numvotes`. Recall from lecture that the syntax for iterating through the key-value pairs of a dictionary `D` is:
``` python
for key, value in D.items():
    ...
```

```python
multivote_voters = set()
# Iterate  over the key-value pairs in voter2numvotes
for ...
	# If the voter has more than one vote in voter2numvotes, add them to multivote_voters.
    ...
```

```python
grader.check("q1p2")
```

## Question 1.3

It turns out that there was only one person who engaged in fraudulent voting. Store their name to `repeat_voter` and the number of times they voted to `num_votes`. 

```python
repeat_voter = ...
num_votes = ...
```

```python
grader.check("q1p3")
```

## Question 1.4

Next we'll determine the winner or winners (in case of a tie) of the election by counting up the votes for each of the candidates. We'll do this in three steps:

1. Create the set of unique candidates (`unique_candidates`)
2. Count the votes for each candidate and store it in a dictionary (`candidate2votes`)
3. Determine the winner(s) by searching for the candidate(s) with the most votes in `candidate2votes`.

**Step 1.** Find the set of unique names of candidates.


```python
unique_candidates = ...
```

```python
grader.check("q1p4p1")
```

**Step 2**. Make a dictionary (`candidate2votes`) that maps each candidate's name to the number of votes that they got.

```python
# 2.1 Initialize the dictionary using dict.fromkeys and setting all values to zero.
candidate2votes = dict.fromkeys(...,...)

# 2.2 Iterate simultaneously through `voters` and `preferred_candidates`
for voter, preferred_candidate in zip(...,...):

    # Skip any voter that voted more than once (ie is present in multivote_voters)
    if ...:
        continue

    # Otherwise, increase the votes for their preferred_candidate by 1.
    candidate2votes[...] ...
```

```python
grader.check("q1p4p2")
```

**Step 3**. Determine the winner(s) by searching for the candidate(s) with the most votes in `candidate2votes`.

```python
# 3.1 Create and empty set called `winners` for storing the winners
winners = ...

# 3.2 The variable `maxvotes` will track the maximum number of votes obtained.
#     Initialize it to 0.
maxvotes = ...

# 3.3 Iterate throught the items in `candidate2votes`
#     Each key is a candidate, and the corresponding value is their votes
for candidate, votes in ...:

    # If this candidate obtained *more* than maxvotes, then it is the winner so far.
    # In this case a) set maxvotes to the new maximum, and b) reset the `winners` set to 
    # a set containing only this candidate
    if ...:
        maxvotes = ...
        winners = ...

    # If the candidate obtained exactly maxvotes, then they should be added to the set of winners. 
    elif ...:
        ...
```

```python
grader.check("q1p4p3")
```

# Question 2: Calculate the bill

In this question you will write a program that calculates the amount owed by each guest at a Vietnamese restaurant. The menu is given below in the form of a dictionary that maps the name of each item to its price in dollars. 

```python
menu = {
    'Crispy Pancake':               11.50,
    'Fresh Spring Rolls':            6.00,
    'Grilled Pork Sausage':          9.00,
    'Lychee Smoothie':               4.50,
    'Shrimp Summer Rolls':           7.00,
    'Turmeric Noodle Soup':          8.00,
    'Chicken Noodle Soup':          10.00,
    'Chicken Rice':                 11.50,
    'Crab Noodle Soup':             13.50,
    'Grilled Pork Vermicelli Bowl': 12.50,
    'Pho Bo':                       10.00,
    'Caramelized Clay Pot Fish':    14.00,
    'Beer':                          8.50,
    'Fresh Coconut Water':           5.00,
    'Jasmine Tea':                   3.50,
    'Vietnamese Iced Tea':           4.50,
    'Vietnamese Iced Coffee':        4.50,
    'Sugarcane Juice':               3.50
}
```

The names of the guests are provided in a list called `guests`. The orders of each of the guests are contained in `orders`, which is a list of lists. The i'th entry of `orders` is the order for the i'th guest. So, for example, Li Wei had fresh spring rolls as an appetizer, pho bo as the main dish, and a Vietnamese iced tea to drink.

```python
guests = [ 'Li Wei',
           'Brian Thompson',
           'Priya Sharma',
           'Sophie Müller',
           'Xiao Mei',
           'Geronimo Marquez' ]

orders = [
    ['Fresh Spring Rolls','Pho Bo','Vietnamese Iced Tea'],
    ['Crispy Pancake','Grilled Pork Vermicelli Bowl','Beer'],
    ['Shrimp Summer Rolls','Chicken Rice','Vietnamese Iced Tea'],
    ['Grilled Pork Sausage','Chicken Noodle Soup','Fresh Coconut Water'],
    ['Crispy Pancake','Crab Noodle Soup','Lychee Smoothie'],
    ['Turmeric Noodle Soup','Chicken Rice','Jasmine Tea']
]
```

The task is to write a function called `compute_bill` that calculates how much each person owes. 

`compute_bill` takes as input the `menu`, `guests`, and `orders` defined above. 

It returns a dictionary (which we call `guest_owes`) that maps each name to their total bill in dollars, and rounded to the nearest cent. 

Each guest leaves an 18\% tip.

**Hint** You can use either `enumerate` or `zip` to iterate simultaneously through `guests` and `orders`.

```python
def compute_bill(menu, guests, orders):
    tip = 0.18
    guest_owes = dict.fromkeys(...,...)

    return guest_owes
```

```python
# Test your code
compute_bill(menu, guests, orders)
```

```python
grader.check("q2p1")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1p1 (points: None)
```
>>> len(voter2numvotes) == len(unique_voters) and sum(voter2numvotes.values()) == len(voters)
True
```
```
>>> isgood = True
>>> for name, numvotes in voter2numvotes.items():
...     isgood &= numvotes == 1 or name == 'Dr Votalot'
>>> isgood and get_hash(voter2numvotes['Dr Votalot']) == '33e75ff09dd601bbe69f351039152189'
True
```


Tests for q1p2 (points: None)
```
>>> len(multivote_voters) == 1
True
```
```
>>> get_hash(len(multivote_voters)) == 'c4ca4238a0b923820dcc509a6f75849b' and get_hash(list(multivote_voters)[0]) == 'b39c6fbe3b1da1f06ee7250178fda15e'
True
```


Tests for q1p3 (points: None)
```
>>> len(repeat_voter) == 10 and num_votes > 10
True
```
```
>>> get_hash(num_votes) == '33e75ff09dd601bbe69f351039152189' and get_hash(repeat_voter) == 'b39c6fbe3b1da1f06ee7250178fda15e'
True
```


Tests for q1p4p1 (points: None)
```
>>> len(unique_candidates) == 7
True
```
```
>>> {get_hash(c) for c in unique_candidates} == {'2bf6c1cbe365029fbd420515acb5d365', '3f158468c312b7d27314c43ba5f54da3', '616ee7ab6ea0db71b0571f690410f674', '7054f8890aa9e7895e5ff73cf3191b3e', '952792a342e27272667f89dd35573b6f', '97a7a1d4b74bf4c5a2ed8ebb2071259b', 'bd0b65b2350f60cfdf6545cf7ecdfeb6'}
True
```


Tests for q1p4p2 (points: None)
```
>>> len(candidate2votes) == 7 and sum(candidate2votes.values()) == 3196
True
```
```
>>> get_hash(candidate2votes['Phaedra Jaka']) == 'a8e864d04c95572d1aece099af852d0a' and get_hash(candidate2votes['Joanne Williem']) == '73278a4a86960eeb576a8fd4c9ec6997' and (get_hash(candidate2votes['Anthony Liebl']) == 'b4288d9c0ec0a1841b3b3728321e7088') and (get_hash(candidate2votes['Rodney Linkoln']) == '115f89503138416a242f40fb7d7f338e') and (get_hash(candidate2votes['Pepe Heydt']) == '310dcbbf4cce62f762a2aaa148d556bd') and (get_hash(candidate2votes['Butler Cardona']) == 'dc6a6489640ca02b0d42dabeb8e46bb7') and (get_hash(candidate2votes['Adriann Cespedes']) == '85fc37b18c57097425b52fc7afbb6969')
True
```


Tests for q1p4p3 (points: None)
```
>>> len(winners) == 1 and maxvotes > 700 and (maxvotes < 800)
True
```
```
>>> {get_hash(c) for c in winners} == {'97a7a1d4b74bf4c5a2ed8ebb2071259b'} and get_hash(maxvotes) == 'a8e864d04c95572d1aece099af852d0a'
True
```


Tests for q2p1 (points: None)
```
>>> g = compute_bill(menu, guests, orders)
>>> isinstance(g, dict) and len(g) == 6
True
```
```
>>> g = compute_bill(menu, guests, orders)
>>> all([v > 20 and v < 40 for v in g.values()])
True
```
```
>>> g = compute_bill(menu, guests, orders)
>>> get_hash(g['Li Wei'], 4) == 'cc53ac578dfb6679a36d00a74e6dac24' and get_hash(g['Brian Thompson'], 4) == '450ac0b261f6b208993e1d8464c57923' and (get_hash(g['Priya Sharma'], 4) == '6e9440c9948ac76d80b739249f9d518d') and (get_hash(g['Sophie Müller'], 4) == '2178f2c3a1286045cd325f29a0137495') and (get_hash(g['Xiao Mei'], 4) == '049487599dfc3da3e35377701157667c') and (get_hash(g['Geronimo Marquez'], 4) == '6e9440c9948ac76d80b739249f9d518d')
True
```
```
>>> g = compute_bill(menu, guests, [['Beer', 'Beer'], ['Fresh Spring Rolls'], ['Shrimp Summer Rolls'], ['Grilled Pork Sausage'], ['Crispy Pancake'], ['Chicken Rice']])
>>> get_hash(g['Li Wei'], 4) == '4daec680d761c4b9b682081edad35fde' and get_hash(g['Brian Thompson'], 4) == 'b5b121f802a17c5b7a2e643530ef561e' and (get_hash(g['Priya Sharma'], 4) == '034b764ea2e856b58be98206f324c49b') and (get_hash(g['Sophie Müller'], 4) == '4ebc9eec7720d57b742c90115c3ee49f') and (get_hash(g['Xiao Mei'], 4) == '8c76ad596d3a0af9b9cba97d57d5c779') and (get_hash(g['Geronimo Marquez'], 4) == '8c76ad596d3a0af9b9cba97d57d5c779')
True
```
