```python
# Initialize Otter
import otter
grader = otter.Notebook("disc01.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Discussion session #01 <br></center></h2>

<h1><center>Submitting to Gradescope<br></center></h1>

---


# 1) Edit mode and command mode

At any time, the notebook can either be in <font color=green>**edit mode**</font> or in <font color=dodgerblue>**command mode**</font>. You can toggle between these two by pressing <kbd>Enter</kbd>  for <font color=green>**edit mode**</font> and <kbd>Esc</kbd> for <font color=dodgerblue>**command mode**</font>. The mode is indicated by the color of the border of the selected cell. In <font color=green>**edit mode**</font> the border is green, while in <font color=dodgerblue>**command mode**</font> the border is blue. 

<font color=green>**Edit mode**</font> is for editing the contents of a cell, while <font color=dodgerblue>**command mode**</font> is for cell-level manipulations, such as creating, deleting, and rearranging entire cells. 

Here is a list of useful keyboard shortcuts to use in <font color=dodgerblue>**command mode**</font>:
* <kbd>A</kbd>: Insert a new cell above
* <kbd>B</kbd>: Insert a new cell below
* <kbd>X</kbd>: Cut the cell
* <kbd>C</kbd>: Copy the cell
* <kbd>V</kbd>: Paste the copied cell
* <kbd>Y</kbd>: Change a cell to a code cell
* <kbd>M</kbd>: Change a cell to a markdown cell

These actions can also be done using toolbar buttons and dropdown menus.

# 2) Executing cells

The easiest way to execute a cell is with a keyboard shortcut.

+ <kbd>Shift</kbd>+<kbd>Enter</kbd> runs the currently highlighted cell(s), and then advances to the next cell. The `▶ Run` button in the toolbar does the same.
+ <kbd>Ctrl</kbd>+<kbd>Enter</kbd> (<kbd>Cmd</kbd>+<kbd>Enter</kbd> on macOS) runs the currently highlighted cells without advancing.

Both of these are also available under the **Cell** menu. 

You can run all of the cells in your notebook by clicking **Run All** in the **Cell** menu. 

It is good to get in that habit of restarting your kernel when running all cells. Do this with **Kernel**->**Restart & Run All**. 

Practice running the next couple cells with <kbd>Shift</kbd>+<kbd>Enter</kbd> and <kbd>Ctrl</kbd>+<kbd>Enter</kbd>. 

```python
import random
print("This is a random number: ", random.random())
```

```python
print("This is another random number: ", random.random())
```

# 3) Autograding and Gradescope

We will be using an automatic grader for lab assignments. Each question in a lab assignment will have an **graded cell** followed by a **testing cell**. You will put gradable code into the **graded cell** as instructed in the lab, and then test your code by executing the **testing cell**. The output of the testing cell will let you know whether you have a mistake in the graded cell. It does this by running your code against a number of  tests.


## Example exercise
Given two variables, `a=1` and `b=2`, compute their sum and assign it to variable `c`. 

+ The correct solution to this exercise is to replace the `...` below with `a+b`.  Try that out, and see that the testing cell marks your answer as correct. 
+ Also try an *incorrect* answer, by replacing the `...` with `-2`. Notice what happens when you run the testing cell.
+ Finally, try setting `c=9` and notice that the testing cell claims that this is *correct*, even though it is clearly wrong! This is because the testing cell contains two types of tests: hidden and not hidden, or open. The open tests are the ones you see in Jupyter Hub. They only test very your answer very broadly. For example, here the test is `c>0`. The open tests are worth zero points. The hidden tests are only evaluated in Gradescope. You will see these results when you upload the zip file to Gradescope. You can submit as many times as you like before the deadline of the assignement.   


#### Graded cell:

```python
a = 1
b = 2
c = ...
```

#### Testing cell:

```python
grader.check("q1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> c > 0
True
```
