```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06B.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session #06-B <br></center></h2>

<h1><center>Matplotlib: 3D plots and Animations<br></center></h1>

---

```python
from resources.hashutils import *
import numpy as np
import matplotlib.pyplot as plt
```

# Question 1: 3D line plot



## Part 1.1: Load the data

The trajectory data for the drone is stored in a 4-column table. The columns, in order, are time, x position, y position, and z position. Each row is a sample measurement of the x, y, and z positions of the drone at a particular time. This matrix is stored in a binary file called 'drone_trajectory.npy', which is stored on your current working directory. Load that file into a NumPy array, and call it `traj`. `traj` should have shape equal to `(75,4)`, meaning that there are 75 consecutive samples. 

**Hint**: For binary NumPy files we use [`np.load`](https://numpy.org/doc/stable/reference/generated/numpy.load.html)

```python
traj = ...
```

```python
grader.check("q1p1")
```

## Part 1.2: Nicknames

It is helpful to define variable names for the columns of `traj`, as we will use them repeatedly when plotting. This saves typing, increases clarity, and reduces the chances fpr bugs. 

Define variables `x`, `y`, and `z` for the corresponding columns in `traj`.

```python
time = ...
x = ...
y = ...
z = ...
```

```python
grader.check("q1p2")
```

## Part 1.3: Highest point

We want to point out in the plot the highest elevation (z value) reached by the drone. Find the index of this value. Call it `maxind`.

```python
maxind = ...
```

```python
grader.check("q1p3")
```

## Part 1.4: Plot

We're now ready to make the plot. Here are the specifications:

+ The figure size is `(8,8)`, and has a single axes. 
+ The trajectory is a solid black line of width 1. The marker is a period `'.'`.
+ Use `ax.scatter` to mark the starting point with a green `'o'` marker of size 80. Label this point as `'start'`.
+ Use `ax.scatter` to mark the point of highest elevation with a red `'o'` marker of size 80.  Label this point as `'highest'`.
+ Add a legend with the start and highest points.

**Try this!**: You can engage an advanced renderer for your plot that allows you to zoom and rotate with your mouse. Do this by adding the ["magic" command](https://ipython.readthedocs.io/en/stable/interactive/magics.html):
```python
%matplotlib widget
```
to the plotting cell. Then run it and drag a point with your mouse.

**Desired result:**

<img src="resources/drone.png" width="400" />

```python
# %matplotlib widget
fig1 = plt.figure(...)
ax1 = fig1.add_subplot(...)
...
```

```python
grader.check("q1p4")
```

# Question 2: Surface plot on a regular grid

The function $f$ shown below produces a real number for each pair of real numbers $(x,y)$.
$$f(x,y) = \cos(5x)\cos(5y)e^{-0.8(x^2 + y^2)}$$
In this part we will use matplotlib's [`plot_surface`](https://matplotlib.org/stable/api/_as_gen/mpl_toolkits.mplot3d.axes3d.Axes3D.plot_surface.html) method to visualize $f(x,y)$ on a square patch of points in $xy$. `plot_surface` take 3 inputs: `X`, `Y`, and `Z`, where `X` and `Y` are mesh matrices obtained from NumPy's [`meshgrid`](https://numpy.org/doc/stable/reference/generated/numpy.meshgrid.html) method. Our first task will be to create the 1D arrays of $x$ and $y$ values required by `meshgrid`.

## Question 2.1: Arrays of x and y values

The grid of points in the $xy$ plane is shown below. To create this, we must first define two 1D arrays with the values of $x$ and $y$. We will call these `xvals` and `yvals`. They should each contain 100 uniformly spaced values ranging from -2 to 2 (inclusive).

<img src="resources/grid.png" width="400" />

**Hint**: [np.linspace](https://numpy.org/doc/stable/reference/generated/numpy.linspace.html)

```python
xvals = ...
yvals = ...
```

```python
grader.check("q2p1")
```

## Question 2.2: Create the mesh

Use [`np.meshgrid`](https://numpy.org/doc/stable/reference/generated/numpy.meshgrid.html) to create the 2D grid of points from the 1D arrays of $x$ and $y$ values. Use variables `X` and `Y` for the output of `meshgrid`.  

```python
X, Y = ...
```

```python
grader.check("q2p2")
```

## Question 2.3: Function evaluation

Evaluate the $f(x,y)$ of the grid. This can be done using NumPy's mathematical functions (such as [`np.exp`](https://numpy.org/doc/stable/reference/generated/numpy.exp.html)) on the grid matrices `X` and `Y`. Store the result as `Z`. `Z` should have shape `(100,100)`. 

```python
Z = ...
```

```python
grader.check("q2p3")
```

## Question 2.4: Surface plot

Next, create a surface plot of the function using [plot_surface](). The plot should have these characteristics:

+ figsize=(8,8)
+ Set the transparency level (`alpha`) of the surface plot to `0.8`. 
+ Set the colormap (`cmap`) of the surface plot to `'BuPu'`.
+ Set the x-axis label to `'x'` with a font size of 20. 
+ Set the y-axis label to `'y'` with a font size of 20. 
+ Try some of the other available colormaps (provided in lecture). Do you have a favorite?

**Desired result:**

<img src="resources/surf.png" width="400" />

```python
fig2 = plt.figure(figsize=...)
ax2 = ...
... # etc.
...
```

```python
grader.check("q2p4")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1p1 (points: None)
```
>>> traj.shape == (75, 4)
True
```
```
>>> get_hash(traj, 4) == '5a0cd86b88b1c12b47c63a6ce3d7c428'
True
```


Tests for q1p2 (points: None)
```
>>> x.shape == (75,) and y.shape == (75,) and (z.shape == (75,))
True
```
```
>>> get_hash(time, 4) == 'd81f6821862a9606640592e2d8f3156f'
True
```
```
>>> get_hash(x, 4) == 'a041a8b0ca5b684d84e70dd51acd5bc5' and get_hash(y, 4) == 'cdd5ce9e443ca575dea0a8af34450d8b' and (get_hash(z, 4) == '289d3a48ce330b5a623e5b0dae5703a7')
True
```


Tests for q1p3 (points: None)
```
>>> maxind > 20 and maxind < 100
True
```
```
>>> get_hash(maxind, 2) == '4a646b1b2110f8fa7a3f4a6cedb151b0'
True
```


Tests for q1p4 (points: None)
```
>>> np.all(np.isclose(fig1.get_size_inches(), (8, 8)))
True
```
```
>>> import mpl_toolkits
>>> isinstance(ax1, mpl_toolkits.mplot3d.axes3d.Axes3D)
True
```
```
>>> len(ax1.get_lines()) == 1
True
```
```
>>> line = ax1.get_lines()[0]
>>> line.get_linestyle() == '-' and line.get_linewidth() == 1 and (line.get_marker() == '.')
True
```
```
>>> len(ax1.collections) == 2
True
```
```
>>> point = list(filter(lambda x: x.get_label() == 'start', ax1.collections))[0]
>>> np.all(np.isclose(point.get_offsets(), np.array([[-0.03093337, -0.02147269]]), atol=0.01))
True
```
```
>>> point = list(filter(lambda x: x.get_label() == 'start', ax1.collections))[0]
>>> np.all(np.isclose(point.get_facecolor(), np.array([[0.0, 0.50196078, 0.0, 1.0]]), atol=0.01))
True
```
```
>>> point = list(filter(lambda x: x.get_label() == 'start', ax1.collections))[0]
>>> point.get_sizes()[0] == 80
True
```
```
>>> point = list(filter(lambda x: x.get_label() == 'highest', ax1.collections))[0]
>>> np.all(np.isclose(point.get_offsets(), np.array([[0.0626858, 0.03630905]]), atol=0.01))
True
```
```
>>> point = list(filter(lambda x: x.get_label() == 'highest', ax1.collections))[0]
>>> np.all(np.isclose(point.get_facecolor(), np.array([[1, 0, 0, 1]]), atol=0.01))
True
```
```
>>> point = list(filter(lambda x: x.get_label() == 'highest', ax1.collections))[0]
>>> point.get_sizes()[0] == 80
True
```
```
>>> set(ax1.get_legend_handles_labels()[1]) == {'highest', 'start'}
True
```


Tests for q2p1 (points: None)
```
>>> xvals.shape == (100,) and xvals[0] == -2 and (xvals[-1] == 2)
True
```
```
>>> yvals.shape == (100,) and yvals[0] == -2 and (yvals[-1] == 2)
True
```
```
>>> get_hash(xvals, 4) == '4f99404ab90fd06b5604744af9af370f' and get_hash(yvals, 4) == '4f99404ab90fd06b5604744af9af370f'
True
```


Tests for q2p2 (points: None)
```
>>> X.shape == (100, 100) and Y.shape == (100, 100) and np.all(X[0, :] == xvals) and np.all(Y[:, 0] == yvals)
True
```


Tests for q2p3 (points: None)
```
>>> Z.shape == (100, 100)
True
```
```
>>> get_hash(Z, 3) == 'ab910c8905e6972bd48338b1effbbb9b'
True
```


Tests for q2p4 (points: None)
```
>>> np.all(np.isclose(fig2.get_size_inches(), (8, 8)))
True
```
```
>>> import mpl_toolkits
>>> isinstance(ax2, mpl_toolkits.mplot3d.axes3d.Axes3D)
True
```
```
>>> ax2.collections[0].get_alpha() == 0.8 and ax2.get_xlabel() == 'x' and (ax2.get_ylabel() == 'y') and (ax2.collections[0].get_cmap().name == 'BuPu')
True
```
