```python
# Initialize Otter
import otter
grader = otter.Notebook("w02Blab.ipynb")
```

```python
import math
from resources.hashutils import get_hash
from resources.utils import is_prime
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session: Week 2-B <br></center></h2>

<h1><center>if statements and while loops<br></center></h1>

---

# Question 1: Counting primes

`utils.py` in the `resources` folder contains a function called `is_prime(n)` that takes as input a positive integer `n` and returns the boolean `True` if that integer is a prime number, and `False` if it is not. Here we will write a new function called `count_primes(n)` that will use `is_prime` internally to count the number of primes that are less than `n`. 

For example, if we call `count_primes(6)`, the function should return the number 3, since there are 3 prime numbers less than 6, namely 2, 3, and 5 (see a table of primes [here](https://upload.wikimedia.org/wikipedia/commons/5/50/Primencomposite0100.svg)).

And if we call `count_primes(7)`, the function should also return 3, since there are only 3 primes **less** than 7 (note that 7 is a prime, but it is not less than 7).

You can use the next cell to play with the `is_prime` function and to work on the `while` loop that will accomplish the goal. 

```python
is_prime(7)
```

Once you are satisfied with your code, you can run the autograder by pasting it into the cell below, between the lines that say "PASTE AND INDENT YOUR CODE HERE" and "DO NOT MODIFY BELOW THIS LINE". 

**Note** You should add one level of indentation to your code so that it is part of the body of the `count_primes` function.

### Autograded cell:

```python
def count_primes(n):
    numprimes=0
    ######### PASTE AND INDENT YOUR CODE HERE ##########

    ######### DO NOT MODIFY BELOW THIS LINE ###########
    return numprimes
```

```python
grader.check("q1")
```

# Question 2: Estimating $\pi$

The convergent series

$$ \pi = 4 - \frac{4}{3} + \frac{4}{5} - \frac{4}{7} + ... $$

known as the *Madhava–Leibniz* series, can be used to estimate $\pi$ to an arbitrary level of precision. The mathematical expression of the series reads as follows:
$$ \pi =  \lim_{n\rightarrow\infty} \sum_{k=0}^{n} a_k$$
where each $a_k$ is given by:
$$ a_k=(-1)^k \frac{4}{2k+1}$$

Write a code that approximates $\pi$ using the Madhava–Leibniz series. It should start by setting the total sum (`mysum`) to zero, and then proceed to add a new term $a_k$ to `mysum` in each iteration of a while loop. It should stop when the absolute value of $a_k$ becomes smaller than the given $\tau$. 

In other words, given $\tau$, the code should return the approximation $\tilde{\pi}$ of $\pi$ given by,

$$\tilde{\pi} = \sum_{k=0}^{K} a_k $$

where $|a_K|<\tau$ and $|a_{K-1}|>\tau$

```python
def estimate_pi(tau):
    mysum = 0
    k = 0
    ######### PASTE AND INDENT YOUR CODE HERE #########
    ...
    ######### DO NOT MODIFY BELOW THIS LINE ###########
    return mysum
```

```python
grader.check("q2")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> count_primes(6) == 3 and count_primes(7) == 3
True
```


Tests for q2 (points: None)
```
>>> math.isclose(estimate_pi(1), 3.466666, abs_tol=0.001)
True
```
```
>>> math.isclose(estimate_pi(0.5), 3.3396825, abs_tol=0.001)
True
```
```
>>> get_hash(estimate_pi(0.1), 4) == '892a68d8b8759e4a97caccc17d27281c'
True
```
```
>>> get_hash(estimate_pi(0.001), 4) == 'cf95c75d1193b553848e157fa690cc97'
True
```
