```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05A.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session #05-A <br></center></h2>

<h1><center>Functions and Modules<br></center></h1>

---

In this lab exercise we will practice making our own modules and running them on our laptops. The submission for this lab is a `.py` file. **You should not submit this notebook, nor does it generate a zip file**. This notebook is meant only as an aid for you to build the `.py` file that you will submit to Gradescope (`mymodule.py`). Once you get through the steps and tests below, you can follow the instructions at the end to assemble `mymodule.py` for final testing and submission to Gradescope.

There are two main reasons to put code into a module, as opposed to a Jupyter notebook. First, a module can be imported and thus reused within other modules or Jupyter notebooks. This is the key to building large and complex programs. Second, again in contrast to Jupyter notebooks, modules can be run as *scripts* from a terminal.

The module that we will build here is very simple (and not very useful). It contains a function called `do_operation(csv,op)` that takes two arguments:

+ `csv`: a comma-separated string of numbers, for example `'1.2,0.5,3.6'`. 
+ `op`: a string that indicates one of three operations to perform on these numbers:`'sum'`, `'prod'`, or `'sort'`. These stand for "summation", "product", and "sort".

`do_operation` returns the result of performing operation on the numbers.

```python
from resources.hashutils import *
import numpy as np
```

# Part 1: Convert the csv string to a NumPy array.

Before implementing `do_operation`, we must first write an auxiliary function called `csv2array(csv)` that takes the csv string as input and returns a NumPy array with the numerical values.  For example, with input `'1.2,0.5,3.6'`, `csv2array` should return the array `np.array([1.2,0.5,3.6])`.

**Notes**: 
+ The dtype of the returned array should be `float64`.
+ Don't worry about checking that the csv string contains only numbers -- the autograder will not test the case that the csv has non-numerical values.
+ Also don't worry about the case that the input is an empty string.
+ A neat challenge is to try to do this in one line of code, using a list comprehension.


```python
def csv2array(csv):
    ...
```

```python
grader.check("q1")
```

# Part 2: Write the `do_operation(csv,op)` function.

As described earlier, this function takes a csv string as input, along with an operation to perform on the numbers in the string. The operation can be either to add them up (`'sum'`), to take their product (`'prod'`), or to sort them (`'sort'`).

Here are additional specifications for `do_operation`:
1. The `'sum'` and `'prod'` operations return a float. The `'sort'` operation returns a NumPy array.
2. If omitted in the input, the operation should default to `'sum'`.

```python
...
```

```python
grader.check("q2")
```

# Part 3: From notebook to module

If you passed the previous tests, you are now ready to convert your code into a Python module that you can run on your laptop. Follow these steps:

1. First make sure that you have Python installed on your computer. You can find instructions for doing this in the document titled "Installing Python", which is stored in bCourses->Files.  

2. Open any text editor.

3. Copy-paste your two functions (`csv2array` and `do_operation`) into the editor. Be sure to put `csv2array` first, since it is used by `do_operation`. Remember to import NumPy at the top of the file.

4. Save this file as `mymodule.py`.

That's it! You can now import and use these two functions into other modules or notebooks. 

# Part 4: Try it out.

This part is optional, but important. We will import `mymodule.py` into a new Jupyter notebook. You may not have a way of running Jupyter notebooks locally on your laptop yet. We will learn how to do that when we learn about VSCode. But you *can* create a new notebook in JupyterHub. 

1. Log in to the UC Berkeley DataHub at [https://data.berkeley.edu/datahub](https://data.berkeley.edu/datahub). Then click on "DATAHUB LOGIN". This will take you to your JupyterHub Home Page. 

2. Upload your `mymodule.py` file with the "Upload" button. Drag-and-drop also works.

3. Create a new notebook by clicking on New->Python 3 (ipykernel).

4. Import the module and use it. Type the following into the first cell of the notebook. Then run it and see that it performs as expected.


<img src="resources/moduletest.png" width="500" />


# Part 5: From module to script

It can sometimes be useful to run python scripts from the terminal (a.k.a. the shell, command prompt, or command line). For example we'd like to be able to type:
```
python mymodule.py prod 4.3,5.4,3.07
```
into the terminal and obtain 71.2854 (the product of 4.3, 5.4, and 3.07). Such a runnable module is called a *script*.

To convert a module into a script that can receive input parameters (the "prod" and "4.3,5.4,3.07" in the example above) from the command line, we need to add a special 'if' clause. 
```python 
if __name__=="__main__":
    ...
```
This should be added to the end of the file, replacing the ellipsis (...) with code to be executed when running the module as a script (as opposed to importing it into another module or notebook). 

In our case we want to pass *two* input arguments from the command line into the script, then pass those along to the `do_operation` function, and print the result. This is accomplished with the following 'if' clause:
```python 
if __name__=="__main__":

    op=sys.argv[1]   # first input argument, e.g. 'prod'
    csv=sys.argv[2]  # second input argument, e.g. '4.3,5.4,3.07'

    value=do_operation(csv,op)
    print(value)
```
Go ahead and add this snippet to the end of `mymodule.py` (the version on your laptop). You will also have to import the `sys` package at the top of `mymodule.py` for this to work. Then test the script in a command prompt (Mac users), PowerShell (Windows users), or terminal (Linux users). Once you are satisfied that it works correctly, you can submit the `mymodule.py` file to Gradescope, where it will be evaluated by the autograder. You can resubmit as many times as you like before the deadline. 

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> A = csv2array('1.1')
>>> isinstance(A, np.ndarray) and A.dtype == 'float64' and (csv2array('1.1,2').shape == (2,))
True
```
```
>>> z = np.random.rand(5)
>>> np.all(np.isclose(csv2array(','.join([str(x) for x in z])), z, rtol=0.0001))
True
```


Tests for q2 (points: None)
```
>>> do_operation('12,2,5', 'invalid command') is None
True
```
```
>>> np.isclose(do_operation('12,2,5', 'sum'), 19) and np.isclose(do_operation('12,2,5', 'prod'), 120)
True
```
```
>>> np.all(np.isclose(do_operation('12,2,5', 'sort'), np.array([2, 5, 12])))
True
```
```
>>> z = np.random.randint(1000, size=5)
>>> randstring = ','.join([str(x) for x in z])
>>> np.isclose(do_operation(randstring), z.sum()).all() and np.isclose(do_operation(randstring, 'prod'), z.prod()).all() and np.isclose(do_operation(randstring, 'sort'), np.sort(z)).all()
True
```
