```python
# Initialize Otter
import otter
grader = otter.Notebook("w06Blab.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session: Week 6-B<br></center></h2>

<h1><center>Modules<br></center></h1>

---

In this lab exercise we will practice making our own modules and running them on our laptops. The submission for this lab is a `.py` file. **You should not submit this notebook, nor does it generate a zip file**. This notebook is meant only as an aid for you to build the `.py` file that you will submit to Gradescope (`mymodule.py`). Once you get through the steps and tests below, you should follow the instructions at the end to assemble `mymodule.py` for final testing and submission to Gradescope.

There are two main reasons to put code into a module, as opposed to a Jupyter notebook. First, a module can be imported and thus reused by other modules and notebooks. This is the key to building large and complex codes. Second, again in contrast to Jupyter notebooks, modules can be run as standalone **scripts** (a.k.a. applications or programs) from the terminal.

The module that we will build here is very simple, and not very useful. It contains a function called `do_operation(op,csv)` that takes two arguments:

+ `op`: a string that indicates one of three operations to perform on these numbers: `'sum'`, `'prod'`, or `'sort'`. These stand for "summation", "product", and "sort".
+ `csv`: a comma-separated string of numbers, for example `'1.2,0.5,3.6'`. 

`do_operation` returns the result of performing the requested operation on the numbers.

```python
from resources.hashutils import *
import numpy as np
```

# Part 1: Convert the csv string to a NumPy array.

Before implementing `do_operation`, we must first write an auxiliary function called `csv2array(csv)` that takes the csv string as input and returns a NumPy array with the numerical values.  For example, with input `'1.2,0.5,3.6'`, `csv2array` should return the array `np.array([1.2,0.5,3.6])`.

**Specification for `csv2array`**
+ The dtype of the returned array should be a float.
+ If `csv` is an empty string, the function should return an array of length 0. 
+ Don't worry about checking that the csv string contains only numbers -- the autograder will not test the case that the csv has non-numerical values.

```python
def csv2array(csv):
    ...
```

```python
grader.check("p1")
```

# Part 2: Write the `do_operation(op,csv)` function.

As described earlier, this function takes a csv string as input, along with an operation to perform on the numbers in the string. The operation can be either to add them up (`'sum'`), to take their product (`'prod'`), or to sort them (`'sort'`).

**Specification for `do_operation`**

1. The `'sum'` and `'prod'` operations return a float. 
2. The `'sort'` operation returns a NumPy array, sorted in ascending order.
3. When `csv` is not passed, it should default to an empty string. 
4. When `op` is not passed, it should default to `'sum'`.
5. `'sum'` applied to an empty string should return 0.0.
6. `'prod'` applied to an empty string should return 0.0.
7. `'sort'` applied to an empty string should return a NumPy array of floats containing no elements.
8. If an invalid operation is passed, the function should return `None`.

```python
...
```

```python
grader.check("p2")
```

# Part 3: From notebook to module

Having completed parts 1 and 2, you are now ready to convert your code into a Python module that you can run on your laptop. Follow these steps:

1. First make sure that you have VSCode and Python installed on your computer (see discussion session for week 5). 
2. Open VSCode.
3. Create a folder for this work. You can call it whatever you like. 
4. Create an environment in the folder using Venv. If you follow the default procedure outlined in the discussion, you will end up with a folder called `.venv` (maybe hidden) in your work folder. 
3. Create a new file called `mymodule.py`.
4. Copy-paste your two functions (`csv2array` and `do_operation`) into the editor. Be sure to put `csv2array` first, since it is used by `do_operation`. Also remember to import NumPy at the top of the file.
5. Now try running the module. It doesn't "do" anything. It just defines two functions. But it will probably fail because we have not added NumPy to the environment. Do that by opening a terminal in VSCode (cmd-j) and typing:
```python
pip install numpy
```
6. Try again to run your module. Again, it doesn't do anything, but it should now not produce any errors. 

# Part 4: Import functionality into a Jupyter notebook. 

Import `mymodule.py` into a new Jupyter notebook. 

1. Create a new file in VSCode called `mynotebook.ipynb`. This file should be in the same folder as `mymodule.py`.
2. Add a first cell to the notebook. Make sure this is a "Python" cell.
3. Import your `do_operation` function into `mynotebook.ipynb` by putting this line into the top cell.

```python
from mymodule import do_operation
```
4. Run the cell. It will ask you to choose an environment; choose your local `.venv` environment. It may ask you to install the `ipykernel` package. Do that as well. This package is needed in any environment where you want to run Jupyter. 

5. If the import works well, you can now use your `do_operation` function in the notebook. 
6. Note that you can also give your function an alias with:

```python
from mymodule import do_operation as doop
```


# Part 5: From module to script.

It can sometimes be useful to run Python scripts from the terminal (a.k.a. the shell, command prompt, or command line). For example we'd like to be able to type:
```
python mymodule.py prod 4.3,5.4,3.07
```
into the terminal and obtain 71.2854 (the product of 4.3, 5.4, and 3.07). Such a runnable module is called a *script*.

To convert a module into a script that can receive input parameters (the "prod" and "4.3,5.4,3.07" in the example above) from the command line, we need to add a special 'if' clause. 
```python 
if __name__=="__main__":
    ...
```
This should be added to the end of the file, replacing the ellipsis (...) with code to be executed when running the module as a script (as opposed to importing it into another module or notebook). 

In our case we want to pass two input arguments from the command line into the script, then pass those along to the `do_operation` function, and print the result. This is accomplished by appending the following code to the end of `mymodule.py`.
```python 
if __name__=="__main__":

    op=sys.argv[1]   # first input argument, e.g. 'prod'
    csv=sys.argv[2]  # second input argument, e.g. '4.3,5.4,3.07'

    value=do_operation(op,csv)
    print(value)
```
Go ahead and add this snippet to the end of `mymodule.py` (the version on your laptop). You will also have to import the `sys` package at the top of `mymodule.py` for this to work. Then test the script in a terminal. 

You can test it in the VSCode by typing, for example, `python mymodule.py prod 4.3,5.4,3.07`.

There may be slight variations to this command depending on your operating system and setup. Ask your GSI for additional help if needed. 

Once you are satisfied that it works correctly, you can submit the `mymodule.py` file to Gradescope, where it will be evaluated by the autograder.

Tests for p1 (points: None)
```
>>> A = csv2array('1.1')
>>> isinstance(A, np.ndarray) and A.dtype == 'float64' and (csv2array('1.1,2').shape == (2,))
True
```
```
>>> len(csv2array('')) == 0
True
```
```
>>> z = np.random.rand(5)
>>> np.all(np.isclose(csv2array(','.join([str(x) for x in z])), z, rtol=0.0001))
np.True_
```


Tests for p2 (points: None)
```
>>> isinstance(do_operation('sum', '1'), float) and isinstance(do_operation('prod', '1'), float)
True
```
```
>>> bool(do_operation(op='sum') == 0.0)
True
```
```
>>> bool(do_operation(op='prod') == 0.0)
True
```
```
>>> bool(do_operation() == 0.0)
True
```
```
>>> np.isclose(do_operation('sum', '12,2,5'), 19) and np.isclose(do_operation('prod', '12,2,5'), 120)
np.True_
```
```
>>> do_operation('invalid operation', '12,2,5') == None
True
```
```
>>> np.all(np.isclose(do_operation('sort', '12,2,5'), np.array([2, 5, 12])))
np.True_
```
```
>>> z = np.random.randint(1000, size=5)
>>> randstring = ','.join([str(x) for x in z])
>>> np.isclose(do_operation(csv=randstring), z.sum()).all() and np.isclose(do_operation('prod', randstring), z.prod()).all() and np.isclose(do_operation('sort', randstring), np.sort(z)).all()
np.True_
```
