```python
# Initialize Otter
import otter
grader = otter.Notebook("w02Alab.ipynb")
```

```python
import math
from resources.hashutils import get_hash
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session: Week 2-A <br></center></h2>

<h1><center>Numerical operations<br></center></h1>

---

# Question 1: Broken phone?

A phone weighing $m=150$ grams is resting on the edge of a table whose surface is $h=75$ centimeters above the floor. The formula for the potential energy of the phone with respect to the floor is:

$$ P  = mgh $$

where $g=9.81\; m/s^2$ is the acceleration due to gravity. 


## Question 1.1

Define variables `m`, `h` and `g` for the mass, gravity, and height respectively. Assign their values in SI units (kilograms, meters, and meters-per-second-squared).

```python
m = ...
h = ...
g = ...
```

```python
grader.check("q1p1")
```

### Question 1.2

Compute the potential energy using the formula above. Assign the result to variable `P`.

```python
P = ...
```

```python
grader.check("q1p2")
```

### Question 1.3

Suddenly the phone falls from the table. Its initial velocity is zero. Let's assume that 2\% of its initial potential energy is lost to friction during the fall, and the remaining 98\% is converted to kinetic energy $K$:

$$ K = 0.98 \times P$$

The kinetic energy is a function of the velocity of the phone as it hits the floor ($v$):

$$ K = \frac{1}{2} m v^2$$

Compute $K$ and $v$.

**Hint**: [`math.sqrt`](https://docs.python.org/3/library/math.html#math.sqrt)

```python
K = ...
v = ...
```

```python
grader.check("q1p3")
```

### Question 1.4

The phone's screen will break whenever it suffers an impact speed of $4 \;m/s$ or greater. Write a logical expression (involving `v`) that evaluates to `True` if the screen breaks and to `False` if it does not. 

```python
screen_breaks = ...
```

```python
grader.check("q1p4")
```

# Question 2: The Great Pyramid of Giza


The Great Pyramid of Giza is located at the base of the Nile delta and is the only of the Wonders of the Ancient World that remains standing. Its height ($h$ in the diagram below) is 146.6 meters, and its base ($b$ in the diagram) is 230.33 meters.

<img src="resources/giza.png" width="800" />

## Question 2.1

Define `h` and `b` as the height and base length of the Great Pyramid in SI units. 

```python
h = ...
b = ...
```

```python
grader.check("q2p1")
```

### Question 2.2

Define `alpha_rad` and `alpha_deg` as the slope of the pyramid ($\alpha$ in the figure) in radians and degrees respectively. 

**Hint**: [math.atan](https://docs.python.org/3/library/math.html#math.atan), [math.degrees](https://docs.python.org/3/library/math.html#math.degrees)

```python
alpha_rad = ...
alpha_deg = ...
```

```python
grader.check("q2p2")
```

### Question 2.3

The volume of a square-based pyramid is given by the formula,
$$ V = \frac{b^2 h}{3} $$
Estimate the volume of the Great Pyramid using this formula, and assign the result to the variable `V`. 

```python
V = ...
```

```python
grader.check("q2p3")
```

### Question 2.4

A friend wishes to build a scale model of the pyramid using a gypsum plaster with a density of $2320\; kg/m^3$. The total weight of the model should be $20\;kg$.

Find the total volume `Vm` of the scale model in $m^3$.

```python
Vm = ...
```

```python
grader.check("q2p4")
```

### Question 2.5

Compute the model's linear scaling factor $s$ as the cube-root of the ratio of `Vm` to `V`:

$ s = \sqrt[3]{\frac{V_m}{V}}$

**Hint** [math.pow](https://docs.python.org/3/library/math.html#math.pow)

```python
s = ...
```

```python
grader.check("q2p5")
```

### Question 2.6

Find the height `hm` of the model (in meters), as well as its base dimension `bm`.

```python
hm = ...
bm = ...
```

```python
grader.check("q2p6")
```

# Question 3: Vibrating rod

A flexible metal rod is rigidly attached to a wall on one side. The other side is struck with a hammer, causing it to vibrate. The displacement of the tip of the rod is measured in meters and denoted with $x(t)$ ($t$ is time in seconds). 

<img src="resources/rod.png" width="700" />

The figure below shows the evolution of $x(t)$ over time. The clock starts at $t=0$ when the rod has reached its maximum positive displacement following the hammer blow ($x_0$). 

<img src="resources/xvt.png" width="1000" />

The formula for $x(t)$ is:

$$x(t)=x_0 \: e^{-at}\cos(\omega t)$$

where $\omega$ is the frequency of oscillations of the rod and $a$ is its damping factor. Their Python variable names and associated numerical values are given in the table below.

| Description    | Math symbol | Python variable     | Value |
| -------- | :-------: | :-------: | ------- |
| initial displacement  | $x_0$    | `x0` | 0.01 meters |
| frequency of oscillations | $\omega$ | `omega` |  200 rad/sec |
| damping factor    | $a$     | `a` | 1 sec $\!{}^{-1}$|

## Question 3.1

Define `x0`, `omega`, and `a`.

```python
x0 = ...
omega = ...
a = ...
```

```python
grader.check("q3p1")
```

## Question 3.2

Compute the value of $x(t)$ after $t=0.07$ seconds. Assign this value to the variable `xA`.

**Hint**: [math.exp](https://docs.python.org/3/library/math.html#math.exp) and [math.cos](https://docs.python.org/3/library/math.html#math.cos)

```python
t = 0.07
xA = ...
```

```python
grader.check("q3p2")
```

## Question 3.3

Using [Euler's formula](https://en.wikipedia.org/wiki/Euler%27s_formula), one can show that the displacement of the rod can also be expressed using complex variables, as follows:

$$x(t) = \text{Re}\left[ \:x_0 \: \exp\left((-a+\omega j)t\right)\: \right]$$

Here "$\text{Re}[z]$" is the real part of a complex variable $z$, and "exp" is the exponential function (ie. powers of Euler's number $e$). 

$j$ is the imaginary unit: $j=\sqrt{-1}$.

Use this formula to compute $x(t)$ at time $t=0.04$. Assign this value to the variable `xB`. 

**Hints**: 
+ [cmath.exp](https://docs.python.org/3/library/cmath.html#cmath.exp)
+ We suggest to compute `xB` in two parts. First define a complex variable `z` as $z =  x_0 \: \exp\left((-a+\omega j)t\right) $, and then obtain its real part using `z.real`.

```python
import cmath
...
xB = ...
```

```python
grader.check("q3p3")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1p1 (points: None)
```
>>> m > 0.1 and m < 0.2 and (h > 0) and (h < 1) and (g > 0) and (g < 10)
True
```
```
>>> get_hash(m, 4) == '9ffc7a03c1959674fa8c0a7369460097'
True
```
```
>>> get_hash(h, 4) == '4c8d9ac7c62b7fbf21fb2f2ae14dd65b'
True
```
```
>>> get_hash(g, 4) == '5c59868ecef993a5d4039e7de4ce5065'
True
```


Tests for q1p2 (points: None)
```
>>> P > 0 and P < 2
True
```
```
>>> get_hash(P, 3) == '777d45bbbcdf50d49c42c70ad7acf5fe'
True
```


Tests for q1p3 (points: None)
```
>>> K > 0 and K < 3 and (v > 0) and (v < 5)
True
```
```
>>> get_hash(K, 4) == '70616f228fd7c3a57a3c3a73c1c1feb6'
True
```
```
>>> get_hash(v, 4) == 'ff3e30920157f452864401806d33a363'
True
```


Tests for q1p4 (points: None)
```
>>> type(screen_breaks) == bool
True
```
```
>>> get_hash(screen_breaks, 1) == '30565a8911a6bb487e3745c0ea3c8224'
True
```


Tests for q2p1 (points: None)
```
>>> h > 100 and h < 150 and (b > 200) and (b < 300)
True
```
```
>>> get_hash(h, 4) == '491ea73dcf0f03080dd30c0c9f4a95c1'
True
```
```
>>> get_hash(b, 4) == '49afb719e0cd46f74578ebf335290f81'
True
```


Tests for q2p2 (points: None)
```
>>> alpha_rad > 0 and alpha_rad < 1 and (alpha_deg > 50) and (alpha_deg < 60)
True
```
```
>>> get_hash(alpha_rad, 4) == '6eb99b694778cc11241bfdeffe698098'
True
```
```
>>> get_hash(alpha_deg, 4) == 'ebac8c3e24e4f84af3d9273e9790535d'
True
```


Tests for q2p3 (points: None)
```
>>> V > 2000000.0 and V < 3000000.0
True
```
```
>>> get_hash(V, 4) == '663e796ebab1f6496f4b74951cb1eb40'
True
```


Tests for q2p4 (points: None)
```
>>> Vm > 0 and Vm < 0.01
True
```
```
>>> get_hash(Vm, 4) == 'c725e129c6bdc0590e2109c0900ec317'
True
```


Tests for q2p5 (points: None)
```
>>> s > 0 and s < 0.02
True
```
```
>>> get_hash(s, 4) == '676da210da0aaaf00daaa94201ccd99b'
True
```


Tests for q2p6 (points: None)
```
>>> hm > 0 and hm < 0.3 and (bm > 0) and (bm < 0.4)
True
```
```
>>> get_hash(hm, 4) == '082a5faddf242a604b6b9118955c5c94'
True
```
```
>>> get_hash(bm, 4) == '443414e5becf86e3be70d28b5e25d976'
True
```


Tests for q3p1 (points: None)
```
>>> x0 > 0 and x0 < 0.1 and (omega > 0) and (omega < 400) and (a > 0) and (a < 3)
True
```
```
>>> get_hash(x0, 4) == '04817efd11c15364a6ec239780038862'
True
```
```
>>> get_hash(omega, 4) == 'bcddc23e127b38859d98f371d3817273'
True
```
```
>>> get_hash(a, 4) == 'e4c2e8edac362acab7123654b9e73432'
True
```


Tests for q3p2 (points: None)
```
>>> xA > 0 and xA < 0.002
True
```
```
>>> get_hash(xA, 4) == 'd80bdf83795e6a20a4f4ee525d595224'
True
```


Tests for q3p3 (points: None)
```
>>> xB < 0 and xB > -0.002
True
```
```
>>> get_hash(xB, 4) == 'e7fc3b2d9a2245fcf3cd75bfaaa1d6ec'
True
```
