---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Discussion week 7 <br></center></h2>

<h1><center>Animations in Matplotlib<br></center></h1>

---

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.animation import FuncAnimation
from IPython.display import HTML
from matplotlib import cm
```

## 1. Adding legends to your plots

+ Must pass `label` argument to each graphical element in the legend.
+ Here we are also adjusting the spines by
    + turning off the top and right spines
    + setting the position of the bottom and left spines to `zero`.

```python
x = np.linspace(0,8,100)
y1 = np.exp(-x)*np.sin(x)
y2 = np.exp(-0.8*x)*np.sin(1.5*x)

fig, ax = plt.subplots(figsize=(8,3))
ax.plot(x,y1,linewidth=2,color='red',linestyle=':')
ax.plot(x,y2,linewidth=2,color='blue',linestyle='--')

ax.set_ylabel('displacement [cm]',fontsize=12)
ax.set_xlabel('time [sec]',loc='right',fontsize=12)
ax.spines[['top','right']].set_visible(False)
ax.spines[['bottom','left']].set_position('zero')
```

## 2. Animating the flight of a drone 

### a) Load the data

```python
traj = np.load('drone_trajectory.npy')
time = traj[:,0]
x = traj[:,1]
y = traj[:,2]
z = traj[:,3]
```

### b) Initialize the figure

+ `ax.plot()` returns a tuple of length 2, and with the line object in the first position. 
+ Setters on the line object include 
    + `set_data_3d`
    + `set_color`
    + `set_linestyle`
    + `set_linewidth`
    + `set_marker`
    + `set_markeredgecolor`
    + `set_markeredgewidth`
    + `set_markerfacecolor`
    + `set_markerfacecoloralt`
    + `set_markersize`
    + `set_alpha`
    + etc.


```python
fig = plt.figure(figsize=(4,4))
ax = fig.add_subplot(projection='3d')
line, = ax.plot(x, y, z, c='black',linewidth=1, linestyle='-')
point, = ax.plot(x[0], y[0], z[0], color='green', marker='o')
```

### c) Figure update function

```python
def update(frame):
    line.set_data_3d(x[:frame+1],y[:frame+1],z[:frame+1])
    point.set_data_3d([x[frame]],[y[frame]],[z[frame]])   
    if frame%2==0:
        point.set_color('red')
    else:
        point.set_color('blue')
```

### d) Create the animation

```python
movie = FuncAnimation(fig=fig, 
                      func=update, 
                      frames=len(x), 
                      interval=100 )
```

### e) Display the animation

```python
HTML(movie.to_html5_video())
```

### e) Save the animation

```python
movie.save('drone_movie.gif')
```

## 3) Animating other graphical elements

### a) y-data of line and bar plots

+ `ax.bar()` returns a collection of the individual `Rectangle` objects.
+ Setters on each `Rectangle` object include 
    + `set_height`
    + `set_linestyle`
    + `set_linewidth`
    + `set_facecolor`
    + `set_edgecolor`
    + `set_alpha`
    + etc.

```python
# data ...............................................
xbar = np.arange(10)
ybar = np.sin(0.5*xbar)

xline = np.linspace(0,10)
yline = np.sin(0.5*xline)

# initialize the figure ..............................
fig, ax = plt.subplots(figsize=(6,3))
bars = ax.bar(xbar,ybar)
line, = ax.plot(xline,yline,linewidth=4,color='purple')
ax.spines[['top','right','bottom']].set_visible(False)
ax.spines['left'].set_position('zero')
ax.set_xticks([])
```

#### Animate the plot so that both line and bar plots move to the right at a rate of one period every 50 frames. That is, 

$$ y = \sin\left(0.5 \:x - \frac{2\pi f}{50} \right) $$

where $f$ is the frame count.

```python
# figure update function .............................
def update(frame):
    ...

# create the animation ...............................
movie = FuncAnimation(fig=fig, 
                      func=update, 
                      frames=50, 
                      interval=100 )

# display the animation...............................
HTML(movie.to_html5_video())
```

### b) Histogram

+ `ax.hist` returns a tuple of length 3 with the histogram object in the third position.
+ The histogram object is a collecton of `Rectangle` objects.

```python
# data ...............................................
data_hist = np.random.normal(loc=0, scale=1, size=100)

# initialize the figure ..............................
fig, ax = plt.subplots(figsize=(6,3))
ax.hist(data_hist,bins=20)
ax.spines[['top','right']].set_visible(False)
```

#### Animate the histogram so that bars are highlighted in order from left to right

```python
# figure update function .............................
def update(frame):
    ...


# create the animation ...............................
movie = FuncAnimation(fig=fig, 
                      func=update, 
                      frames=20, 
                      interval=200 )

# display the animation...............................
HTML(movie.to_html5_video())
```

### b) Scatter

+ `ax.scatter` returns the scatter object.
+ Setters on the scatter object include 
    + `set_facecolor`
    + `set_edgecolor`
    + `set_sizes`
    + etc.

```python
# data ...............................................
x_scatter = np.sort( np.random.rand(20)*10 )
y_scatter = 3*x_scatter**2 + np.random.normal(0,20,20)

# initialize the figure ..............................
fig, ax = plt.subplots(figsize=(6,3))
scatter = ax.scatter(x_scatter,y_scatter,c='black')
ax.spines[['top','right']].set_visible(False)
ax.spines[['bottom','left']].set_position('zero')
ax.grid()
```

#### Animate the scatter plot so that points are highlighted in order from left to right

```python

# figure update function .............................
def update(frame):
    ...

# create the animation ...............................
movie = FuncAnimation(fig=fig, 
                      func=update, 
                      frames=20, 
                      interval=200 )

# display the animation...............................
HTML(movie.to_html5_video())
```

### c) Surface plot

+ `ax.plot_surface` returns the surface object.
+ It is difficult to dynamically change the data in a surface plot. 
+ It is always possible to clear and remake the entire plot in the update function. This is slower but still valid.  
+ Use `ax.clear()` to clear the axes. 

```python
# data ...............................................
x = np.arange(-3,3,0.25)
y = np.arange(-3,3,0.25)
X, Y = np.meshgrid(x, y)
Z = np.cos(np.sqrt(X**2 + Y**2)) + 1

# initialize the figure ..............................
fig = plt.figure(figsize=(4,4))
ax = fig.add_subplot(projection='3d')
ax.plot_surface(X, Y, Z, cmap=cm.RdYlBu,vmin=-1, vmax=1)
ax.set_zlim(-1,1)
```

#### Animate the surface plot to show this time-varying function

$$ f(x,y,t) = \left( \cos\left(\sqrt{x^2 + y^2}\right)  + 1\right)  \cos\left(\frac{2\pi}{50}\:t\right)  $$

```python
# figure update function .............................
def update(frame):
    ...

# create the animation ...............................
movie = FuncAnimation(fig=fig, 
                      func=update, 
                      frames=50, 
                      interval=30 )

# display the animation...............................
HTML(movie.to_html5_video())
```

Tests for q1.0 (points: 1)
```
>>> assert get_hash(a) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
>>> assert get_hash(b) == '596a3d04481816330f07e4f97510c28f'
>>> assert get_hash(c) == 'a87ff679a2f3e71d9181a67b7542122c'
>>> assert get_hash(x) == 'c81e728d9d4c2f636f067f89cc14862c'
```


Tests for q1.1 (points: 1)
```
>>> assert get_hash(round(q1_1, 2)) == '10f8e19d19bf3772687b5457a689fa21'
```


Tests for q1.2 (points: 1)
```
>>> assert get_hash(q1_2) == 'eac1be4ceede350c62d93c8212e3e536'
```


Tests for q1.3 (points: 1)
```
>>> assert get_hash(round(q1_3, 3)) == '5e295664d633f88870fc93eda8cf42c9'
```


Tests for q1.4 (points: 1)
```
>>> assert get_hash(round(q1_4, 3)) == '14a742b8cccaf65017719898cf735d84'
```


Tests for q1.5 (points: 1)
```
>>> assert get_hash(round(q1_5, 4)) == '88491693334ee557168464e310349583'
```


Tests for q1.6 (points: 1)
```
>>> assert get_hash(round(np.real(q1_6),8)) == 'c54490d3480079138c8c027a87a366e3'
>>> assert get_hash(round(np.imag(q1_6),8)) == 'e4c2e8edac362acab7123654b9e73432'
```


Tests for q1.7 (points: 1)
```
>>> assert get_hash(round(q1_7,4)) == '3205c5a0e0c890e0113e87dd3b5d7349'
```


Tests for q1.8 (points: 1)
```
>>> assert (get_hash(round(q1_8, 3)) == 'c0857c1de732e77e7f99af3055dda037')
```


Tests for q2.0 (points: 0.5)
```
>>> assert len(q2_0.upper()) == 1 , 'Make sure you have just the letter of your answer choice in quotes'
>>> assert get_hash(q2_0.upper()) == '9d5ed678fe57bcca610140957afab571'
```


Tests for q2.1 (points: 1.5)
```
>>> var_names = ['China_pop' , 'China_CO2' , 'Qatar_pop' , 'Qatar_CO2' , 'USA_pop' , 'USA_CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(China_pop) == '0e4e946668cf2afc4299b462b812caca'
>>> assert get_hash(China_CO2) == 'b7852f3d3a775b5188139a2c07024397'
>>> assert get_hash(Qatar_pop) == '23cb05d4a2e5bf28363909d92cd244e2'
>>> assert get_hash(Qatar_CO2) == '26657d5ff9020d2abefe558796b99584'
>>> assert get_hash(USA_pop) == 'c042f4db68f23406c6cecf84a7ebb0fe'
>>> assert get_hash(USA_CO2) == '351869bde8b9d6ad1e3090bd173f600d'
```


Tests for q2.2 (points: 1)
```
>>> assert get_hash(type(q2_2)) == '71f335aebe2d722b3351e4c9710a7236' , 'Check which function you used.'
>>> assert get_hash(q2_2) == '2203e3c2293eb8c9bda5bcd97acbf884' , 'Does the value have a fractional component?'
```


Tests for q2.3 (points: 1)
```
>>> assert get_hash(q2_3) == 'c81e728d9d4c2f636f067f89cc14862c'
>>> assert get_hash(Qatar_pop) == '23cb05d4a2e5bf28363909d92cd244e2' , 'Do not reassign the variable'
```


Tests for q2.4 (points: 1)
```
>>> assert get_hash(q2_4) == '2993d6c487b8101c0701534ba325cb91'
>>> assert get_hash(Qatar_CO2) == '26657d5ff9020d2abefe558796b99584' , 'Do not reassign the variable'
```


Tests for q3.0 (points: 1)
```
>>> q3_0_vars = ['China_per_capita' , 'Qatar_per_capita' , 'USA_per_capita']
>>> assert all([var_name in globals() for var_name in q3_0_vars]), 'Check variable names' # Checks that the variables exist
>>> assert get_hash(China_per_capita) == '2e08fabaf262a0fea4da5c818c9894e2' , 'Check variable value'
>>> assert get_hash(Qatar_per_capita) == '2f56f1561f0c21041643671d4fd9f85a' , 'Check variable value'
>>> assert get_hash(USA_per_capita) == 'c89b4e5c27b4928beff41e7782689f91' , 'Check variable value'
```


Tests for q3.1 (points: 1)
```
>>> assert get_hash(q3_1) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q3.2 (points: 1)
```
>>> assert get_hash(q3_2a) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_2b) == 'f8320b26d30ab433c5a54546d21f414c'
>>> assert get_hash(q3_2c) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q3.3 (points: 1)
```
>>> assert get_hash(q3_3a) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_3b) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_3c) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_3d) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q3.4 (points: 1)
```
>>> assert get_hash(round(q3_4a,3)) == '302ddfd6d63bbe5693b96981c129d3e2'
>>> assert get_hash(round(q3_4b,3)) == '70e3b658922334c673fa2959dbfdfbd6'
>>> assert get_hash(q3_4c) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q3.5 (points: 1)
```
>>> assert get_hash(q3_5) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q4.0 (points: 1)
```
>>> var_names = ['pop' , 'CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(pop_q4_0) == '5e784975e5eb0ff3a28a2aad3953dcbb'
>>> assert get_hash(CO2_q4_0) == '8f484d873ab271148f32f7d0921afb08'
```


Tests for q4.1 (points: 1)
```
>>> var_names = ['pop' , 'CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(pop_q4_1) == 'd0ecac9d532783c816b4995f1327a2e6'
>>> assert get_hash(CO2_q4_1) == '8f484d873ab271148f32f7d0921afb08'
```


Tests for q4.2 (points: 1)
```
>>> var_names = ['pop' , 'CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(pop) == '881602830957a3cd082095af4601bb6e'
>>> assert get_hash(CO2) == 'e88a56268abff0794d40bf0f030ed684'
```


Tests for q4.3 (points: 1)
```
>>> var_names = ['per_capita']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(per_capita) == '7b6ee17515b452692f05c2c2341c91cf'
```


Tests for q4.4 (points: 1)
```
>>> assert get_hash(q4_4) == '977c8f03a4a2c1813c6bcaaa0b831963'
```
