```python
# Initialize Otter
import otter
grader = otter.Notebook("lab09A.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session #09-A <br></center></h2>

<h1><center>Bisection and Newton-Raphson for solving nonlinear equations<br></center></h1>

---

```python
import numpy as np
import matplotlib.pyplot as plt
from resources.hashutils import *
```

# Question 1: Spherical container

<center><img src="resources/sphere.png" width=300/></center>

The figure above shows a spherical glass container used to mix liquids. The container has a hole through which the liquids are poured, and a small propeller to do the mixing. Unfortunately the container is old and the tick marks indicating the volume of liquid are no longer legible. The only way to know the amount of liquid in the container is to measure the height $h$ of the fluid's surface, and then to use the formula for the volume of a partially filled sphere:

$$ V = \frac{\pi h^2}{3} \left(3r - h \right) $$

Here $r=10$ cm is the radius of the sphere and $V$ is the volume in cubic centimeters (i.e. milliliters)

```python
r = 10   # radius in centimeters
```

## Question 1.1: Volume evaluation

Write a function called `evalV` that takes the fluid level `h` (in centimeters), and returns to volume of fluid in milliliters, according to the formula just given. 

**Note**
+ `r` is considered a global constant and therefore not passed through the arguments. 

```python
def evalV(h):
    return ...
```

```python
grader.check("q1p1")
```

## Question 1.2: Total volume

What is the total volume of the sphere in liters?

```python
total_volume = ...
```

```python
grader.check("q1p2")
```


A chemical recipe requires that you mix 3 liters of a liquid using the glass sphere. You decide to use the root-finding methods learned in class to find the level $h$ corresponding to this volume of liquid. You begin by plotting the function $V(h)$ in order to get a sense of the answer.

## Question 1.3: Plot volume versus height

Make the plot shown in the figure. The plot shows
+ the volume of liquid as a function of height $h$ (magenta line), and
+ the desired volume (dotted blue line)

Other specifications:

+ The x-axis grid for the magenta line should consist of fifty values ranging from 0 to $2r$ (inclusive).
+ `figsize=(8,4)`
+ The grid and legend as shown in the plot.

<img src="resources/volvsheight.png" width=400 />

**Note**
+ Plot the horizontal line using `plot` or [`axhline`](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.axhline.html), but **not hlines**, otherwise the autograder will fail. 


```python
fig1, ax1 = plt.subplots(...)
...
```

```python
grader.check("q1p3")
```

## Question 1.4: Equation to solve

We can see in the figure above that 3 liters of fluid corresponds to a height of a little more than 12.5 centimeters. To solve this more precisely, we must find the exact point at which the function $V(h)$ crosses the horizontal line at 3000 milliliters. In other words, we must find the roots of the function $f(h)$:

$$ f(h)=V(h)-3000$$

Define the function `evalf(h)` that evaluates $f(h)$.

```python
def evalf(h):
    return ...
```

```python
grader.check("q1p4")
```

## Question 1.5: Number of bisection steps

Write a function called `compute_bisection_steps(N0,P0,maxwidth)` that computes the number of bisection steps needed to reduce the interval width from `|P0-N0|` to less or equal than `maxwidth`. The formula for doing this was derived in lecture and is reproduced below.

$$ k = \text{ceil} \left( \log_2 \left( \frac{ |P_0-N_0| }{ \bar{w}} \right)  \right) $$

**Hint**:
+ NumPy's [ceil](https://numpy.org/doc/stable/reference/generated/numpy.ceil.html) method returns a float. You should cast this to an integer in order to avoid problems in later parts.
+ Return `0` whenever $|P_0-N_0|<\bar{w}$. 

```python
def compute_bisection_steps(N0,P0,maxwidth):
    ...
```

```python
# You can use this cell to test your code
maxwidth = 1
N0 = 0
P0 = 0.6
numsteps =  compute_bisection_steps(N0,P0,maxwidth)
numsteps
```

```python
grader.check("q1p5")
```

## Question 1.6: Bisection

Implement the bisection algorithm as it was described in lecture. 

Carefully follow the comments in the template. 

The `bisection` method takes as inputs the initial interval values `N0` and `P0`, and the desired `maxwidth`. Here, `N0` is the value of $h$ corresponding to a *negative* evaluation of $f(h)$, and `P0` is the value of $h$ corresponding to a *positive* evaluation of $f(h)$. The `bisection` method should return two NumPy arrays, `N` and `P`, corresponding to the the *negative* and *positive* sequences.

**Note**
+ In the case where $|P_0-N_0|<\bar{w}$, the code should return `(np.array(N0), np.array(P0))`.



```python
def bisection(N0,P0,maxwidth):

    # Use the method defined in a previous part to compute 
    # the number of bisection steps to execute.
    numsteps =  ...

    # If no steps are taken, return the initial guesses
    if numsteps==0:
        return ..., ...

    # Initialize the negative and positive sequences
    # These are 1D NumPy arrays of length numsteps
    # They will be filled completely, so you can use np.empty
    # Put N0 and P0 in the first position of N and P respectively
    N = ...
    P = ...
    N[0] = ...
    P[0] = ...

    # k iterates from 1 to numsteps-1
    for k in range(1,numsteps):

        # Find midpoint of the negative and positive sequences
        C = ...
        
        # evaluate f(C)
        fC = ...

        # compute N[k] and P[k]
        ...
        ...

    return N, P
```

```python
# You can use this cell to test your code
maxwidth = 0.1
N0 = 0
P0 = 20
N, P = bisection(N0,P0,maxwidth)
```

```python
grader.check("q1p6")
```

## Question 1.7: Plot

Plot the sequences `N` and `P` computed by the bisection method with parameters `N0=0`, `P0=20`, and `maxwidth=0.1`.
Define `hsol` as the midpoint of the last interval returned by the bisection iteration. 

Other specifications:

+ The `N` and `P` lines have `linewidth=2` and `marker='o'`.  
+ Draw a horizontal line at the level of `hsol`. Again, use `axhline` or `plot`.
+ Include the grid and legend as shown in the plot. 

<center><img src="resources/bisection.png" width=600/></center>



```python
N, P = ...
hsol = ...

fig2, ax2 = plt.subplots(figsize=(8,4))
...
```

```python
grader.check("q1p7")
```

## Question 1.8: Derivative of $f(h)$


Next we will apply the Newton-Raphson algorithm to the same problem. This is possible since the function that we are solving is differentiable, and the derivative can be easily evaluated in code. Implement the function `evalfprime` that returns $f'(h)$: the derivative of $f$ for a given value of $h$. 

Recall:

\begin{align*} 
&V(h) = \frac{\pi h^2}{3} \left(3r - h \right) \\
&f(h)=V(h)-3000
\end{align*}



```python
def evalfprime(h):
    return ...
```

```python
grader.check("q1p8")
```

## Question 1.9: Newton-Raphson

Implement the Newton-Raphson algorithm that was described in lecture.

The function `newton_raphson(x0,taux,tauf,maxsteps)` takes an initial guess `x0`, tolerances `taux` and `tauf`, and the maximum number of steps `maxsteps`. It returns a NumPy array with the sequence of values, beginning with `x0`.

**Note**. 

1. You will notice that the steps do not exactly follow the pseudo-code. The psuedo-code involves a "repeat-until", which in Python could be implemented with a "while" loop. However we do it here with a "for" loop. The reason for this is that, when using NumPy it is better to pre-allocate the entire array (with np.full, np.empty, np.zeros, etc.) than to extend the array in each iteration. But once we allocated a fixed-size array, it is more natural (more explicit, more clear) to iterate through it with a "for" loop than with a "while" loop.

2. Note also that the function is cubic and therefore may have up to three real roots. In fact, this function *does* have three real roots, but only one of them is in the acceptable range of heights, between 0 and 20 cm. This is ok. You should allow the Newton-Raphson method to return these "unphysical" solutions if it finds them.

```python
def newton_raphson(x0,taux,tauf,maxsteps):

    # Allocate x by creating an array of np.nan values of length maxsteps
    x = ...

    # Put the initial guess in the first position of x
    ...

    # Iterate through the un-assigned values in x.
    for k in ...:

        # Find the next value in the sequence according to the Newton-Raphson update formula
        ...
                
        # Break if the difference between the latest two values in the sequence is less than taux
        if ...:
            break

        # Break if the function evaluation is within tauf of zero.
        if ...:
            break

    # If either of the two break conditions were hit, then there will be trailing NaNs in x. 
    # Remove any trailing NaNs. 
    ...

    return x
```

```python
# Use this cell to test your code
x = newton_raphson(x0=20.5,taux=.01,tauf=2,maxsteps=100)
x
```

```python
grader.check("q1p9")
```

## Question 1.10: Compare Bisection and Newton-Raphson

We can now test both methods to see which converges in fewer steps. Note that this does not imply that it takes less *time*, since the steps of one algorithm may be more computationally expensive (and therefore *slower*) than the steps of another algorithm. Between bisection and Newton-Raphson, it would seem that a Newton-Raphson step is more expensive, since it involves evaluations of both $f(h)$ and $f'(h)$, whereas bisection only evaluates $f(h)$.

Make the plot shown below that overlays the progression of both algorithms. Here is the specification. 

For bisection:
+ Run the algorithm with `N0=0`, `P0=20`, and `maxwidth=0.1`.
+ Plot the upper and lower limits of the interval with solid lines and 'o' markers. 
+ Use the [ax.fill_between](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.fill_between.html) method to shade in the space between the two limits (I like to use a transparency value of `alpha=0.2` for this). 
+ Label the "fill_between" element (not the upper and lower lines) as "Bisection".

For Newton-Raphson:
+ Run the algorithm with `x0=5`, `taux=0.1`, `tauf=0`, `maxsteps=10`.
+ Plot the sequence with a solid line and 'o' markers. 
+ Label it as "Newton-Raphson".

For the solution.
+ Consider the last value of the Newton-Raphson sequence to be the solution. 
+ Plot a horizontal dashed line that goes through the solution. 
+ Label it as "solution".

```python
...
fig3, ax3 = plt.subplots(figsize=(8,4))
...
```

```python
grader.check("q1p10")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1p1 (points: None)
```
>>> np.isclose(evalV(1), 30.36872, atol=0.01)
True
```
```
>>> get_hash([evalV(h) for h in np.arange(1, 6)], 5) == 'cab0173151fc8d51e6d2fb7b2c16962f'
True
```


Tests for q1p10 (points: None)
```
>>> z = ax3.get_legend_handles_labels()
>>> fb = z[0][z[1].index('Bisection')]
>>> path = fb.get_paths()[0]
>>> get_hash(np.sort(np.unique(path.vertices[:, 1])), 4) == '469c30ad3f9150bd76272e5cc09b6616'
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax3.get_lines()}
>>> get_hash(plotdata['solution'][0], 4) == 'e5f53791e854c3b33cf13e30406312a2'
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax3.get_lines()}
>>> get_hash(plotdata['Newton-Raphson'], 4) == 'a619bcae08752f158de2e2d07ea47d24'
True
```
```
>>> {line.get_marker() for line in ax3.get_lines() if line.get_label() != 'solution'} == {'o'} and {line.get_linestyle() for line in ax3.get_lines() if line.get_label() != 'solution'} == {'-'}
True
```


Tests for q1p2 (points: None)
```
>>> total_volume > 0 and total_volume < 10
True
```
```
>>> get_hash(total_volume, 4) == 'fd4c55c738a8e8c110920d67f21e93b9'
True
```


Tests for q1p3 (points: None)
```
>>> import matplotlib
>>> isinstance(ax1, matplotlib.axes._axes.Axes)
True
```
```
>>> np.all(fig1.get_size_inches() == [8.0, 4.0])
True
```
```
>>> {line.get_label().lower() for line in ax1.get_lines()} == {'desired volume', 'volume vs. height'}
True
```
```
>>> line1 = list(filter(lambda x: x.get_label() == 'volume vs. height', ax1.get_lines()))[0]
>>> xdata = line1.get_xdata()
>>> xdata.shape == (50,)
True
```
```
>>> line1 = list(filter(lambda x: x.get_label() == 'volume vs. height', ax1.get_lines()))[0]
>>> xdata = line1.get_xdata()
>>> get_hash(xdata, 3) == 'c97958d82c12e4f109e317f1667e1757'
True
```
```
>>> line2 = list(filter(lambda x: x.get_label() == 'desired volume', ax1.get_lines()))[0]
>>> ydata = np.array(line2.get_ydata())
>>> np.all(ydata == 3000)
True
```
```
>>> ax1.get_xgridlines()[0].get_visible()
True
```
```
>>> ax1.get_legend().get_visible()
True
```


Tests for q1p4 (points: None)
```
>>> evalf(0) == -3000 and evalf(20) > 1000
True
```
```
>>> get_hash([evalf(h) for h in range(10, 16)], 4) == '94436faad921ede2ea6ec58b2e554f4b'
True
```


Tests for q1p5 (points: None)
```
>>> type(compute_bisection_steps(0, 20, 1)) == int
True
```
```
>>> compute_bisection_steps(0, 0.4, 1) == 0
True
```
```
>>> [compute_bisection_steps(0, p, 1) for p in np.arange(20, 60, 4)] == [5, 5, 5, 5, 6, 6, 6, 6, 6, 6]
True
```
```
>>> np.all([compute_bisection_steps(0, p, 1) for p in np.arange(20, 100, 7)] == [5, 5, 6, 6, 6, 6, 6, 7, 7, 7, 7, 7])
True
```


Tests for q1p6 (points: None)
```
>>> (N, P) = bisection(0, 10, 3)
>>> np.all(N == np.array([0.0, 5.0])) and np.all(P == np.array([10.0, 10.0]))
True
```
```
>>> (N, P) = bisection(0, 0.3, 0.5)
>>> N == np.array(0) and P == np.array(0.3)
True
```
```
>>> (N, P) = bisection(0, 20, 0.7)
>>> get_hash(N, 3) == '594dbad242cd23ad0a41a496054c9865' and get_hash(P, 3) == '5ea9d08c67c6e7d5d30c2da286174253'
True
```


Tests for q1p7 (points: None)
```
>>> len(ax2.get_lines()) == 3
True
```
```
>>> {line.get_label() for line in ax2.get_lines()} == {'N', 'P', 'hsol'}
True
```
```
>>> xdata = [line.get_xdata() for line in ax2.get_lines() if line.get_label() == 'N'][0]
>>> np.all(xdata == np.arange(8))
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> get_hash(plotdata['hsol'][0], 4) == '9ff0e5612f7da2741e689c26b276a8c2'
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> get_hash(plotdata['N'][0], 4) == '30565a8911a6bb487e3745c0ea3c8224' and get_hash(plotdata['P'][0], 4) == '75cf3ac5e70c76583be3efb5012bd44e'
True
```
```
>>> plotdata = {line.get_label(): line for line in ax2.get_lines()}
>>> plotdata['N'].get_linewidth() == 2.0 and plotdata['N'].get_marker() == 'o' and (plotdata['P'].get_linewidth() == 2.0) and (plotdata['P'].get_marker() == 'o')
True
```


Tests for q1p8 (points: None)
```
>>> evalfprime(10) > 300 and evalfprime(10) < 320
True
```
```
>>> get_hash([evalfprime(h) for h in np.linspace(10, 20, 5)], 5) == '9e45fa83f5b94d7392168566c0fd34cb'
True
```


Tests for q1p9 (points: None)
```
>>> x = newton_raphson(x0=10, taux=1, tauf=1, maxsteps=10)
>>> x[0] == 10
True
```
```
>>> x = newton_raphson(x0=10, taux=1, tauf=1, maxsteps=10)
>>> len(x) == 3
True
```
```
>>> x = newton_raphson(x0=1, taux=1, tauf=1, maxsteps=5)
>>> len(x) == 5
True
```
```
>>> x = newton_raphson(x0=1, taux=0.1, tauf=0.1, maxsteps=100)
>>> get_hash(x, 4) == 'cc67c66000938d314e4801d7ab53358e'
True
```
```
>>> x = newton_raphson(x0=-10, taux=0.1, tauf=1, maxsteps=10)
>>> get_hash(x, 4) == 'dcc0ecf5a677a9fe6a8bfbe0332ec63f'
True
```
