---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Discussion week 5<br></center></h2>

<h1><center>VSCode<br></center></h1>

---

+ [VSCode download page](https://code.visualstudio.com/Download)
+ [Python download page](https://www.python.org/downloads/)
+ Populate your Python environment with:
```python
pip install numpy matplotlib scipy otter-grader jupyter ipykernel
```

or

```python
pip install -r requirements_e7.txt
```


### Debgugging example

The function `find_max` is supposed to find the largest of the values in `numbers`. 

```python
def find_max(numbers):
    max_value = 0
    
    for n in numbers:
        if n > max_value:
            print(f"n={n} max_value={max_value}")
            max_value = n
   
    return max_value
```

```python
themax = find_max([5,3,7])

print( f"Maximum: {themax}" )
```

Tests for q1.0 (points: 1)
```
>>> assert get_hash(a) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
>>> assert get_hash(b) == '596a3d04481816330f07e4f97510c28f'
>>> assert get_hash(c) == 'a87ff679a2f3e71d9181a67b7542122c'
>>> assert get_hash(x) == 'c81e728d9d4c2f636f067f89cc14862c'
```


Tests for q1.1 (points: 1)
```
>>> assert get_hash(round(q1_1, 2)) == '10f8e19d19bf3772687b5457a689fa21'
```


Tests for q1.2 (points: 1)
```
>>> assert get_hash(q1_2) == 'eac1be4ceede350c62d93c8212e3e536'
```


Tests for q1.3 (points: 1)
```
>>> assert get_hash(round(q1_3, 3)) == '5e295664d633f88870fc93eda8cf42c9'
```


Tests for q1.4 (points: 1)
```
>>> assert get_hash(round(q1_4, 3)) == '14a742b8cccaf65017719898cf735d84'
```


Tests for q1.5 (points: 1)
```
>>> assert get_hash(round(q1_5, 4)) == '88491693334ee557168464e310349583'
```


Tests for q1.6 (points: 1)
```
>>> assert get_hash(round(np.real(q1_6),8)) == 'c54490d3480079138c8c027a87a366e3'
>>> assert get_hash(round(np.imag(q1_6),8)) == 'e4c2e8edac362acab7123654b9e73432'
```


Tests for q1.7 (points: 1)
```
>>> assert get_hash(round(q1_7,4)) == '3205c5a0e0c890e0113e87dd3b5d7349'
```


Tests for q1.8 (points: 1)
```
>>> assert (get_hash(round(q1_8, 3)) == 'c0857c1de732e77e7f99af3055dda037')
```


Tests for q2.0 (points: 0.5)
```
>>> assert len(q2_0.upper()) == 1 , 'Make sure you have just the letter of your answer choice in quotes'
>>> assert get_hash(q2_0.upper()) == '9d5ed678fe57bcca610140957afab571'
```


Tests for q2.1 (points: 1.5)
```
>>> var_names = ['China_pop' , 'China_CO2' , 'Qatar_pop' , 'Qatar_CO2' , 'USA_pop' , 'USA_CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(China_pop) == '0e4e946668cf2afc4299b462b812caca'
>>> assert get_hash(China_CO2) == 'b7852f3d3a775b5188139a2c07024397'
>>> assert get_hash(Qatar_pop) == '23cb05d4a2e5bf28363909d92cd244e2'
>>> assert get_hash(Qatar_CO2) == '26657d5ff9020d2abefe558796b99584'
>>> assert get_hash(USA_pop) == 'c042f4db68f23406c6cecf84a7ebb0fe'
>>> assert get_hash(USA_CO2) == '351869bde8b9d6ad1e3090bd173f600d'
```


Tests for q2.2 (points: 1)
```
>>> assert get_hash(type(q2_2)) == '71f335aebe2d722b3351e4c9710a7236' , 'Check which function you used.'
>>> assert get_hash(q2_2) == '2203e3c2293eb8c9bda5bcd97acbf884' , 'Does the value have a fractional component?'
```


Tests for q2.3 (points: 1)
```
>>> assert get_hash(q2_3) == 'c81e728d9d4c2f636f067f89cc14862c'
>>> assert get_hash(Qatar_pop) == '23cb05d4a2e5bf28363909d92cd244e2' , 'Do not reassign the variable'
```


Tests for q2.4 (points: 1)
```
>>> assert get_hash(q2_4) == '2993d6c487b8101c0701534ba325cb91'
>>> assert get_hash(Qatar_CO2) == '26657d5ff9020d2abefe558796b99584' , 'Do not reassign the variable'
```


Tests for q3.0 (points: 1)
```
>>> q3_0_vars = ['China_per_capita' , 'Qatar_per_capita' , 'USA_per_capita']
>>> assert all([var_name in globals() for var_name in q3_0_vars]), 'Check variable names' # Checks that the variables exist
>>> assert get_hash(China_per_capita) == '2e08fabaf262a0fea4da5c818c9894e2' , 'Check variable value'
>>> assert get_hash(Qatar_per_capita) == '2f56f1561f0c21041643671d4fd9f85a' , 'Check variable value'
>>> assert get_hash(USA_per_capita) == 'c89b4e5c27b4928beff41e7782689f91' , 'Check variable value'
```


Tests for q3.1 (points: 1)
```
>>> assert get_hash(q3_1) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q3.2 (points: 1)
```
>>> assert get_hash(q3_2a) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_2b) == 'f8320b26d30ab433c5a54546d21f414c'
>>> assert get_hash(q3_2c) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q3.3 (points: 1)
```
>>> assert get_hash(q3_3a) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_3b) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_3c) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_3d) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q3.4 (points: 1)
```
>>> assert get_hash(round(q3_4a,3)) == '302ddfd6d63bbe5693b96981c129d3e2'
>>> assert get_hash(round(q3_4b,3)) == '70e3b658922334c673fa2959dbfdfbd6'
>>> assert get_hash(q3_4c) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q3.5 (points: 1)
```
>>> assert get_hash(q3_5) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q4.0 (points: 1)
```
>>> var_names = ['pop' , 'CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(pop_q4_0) == '5e784975e5eb0ff3a28a2aad3953dcbb'
>>> assert get_hash(CO2_q4_0) == '8f484d873ab271148f32f7d0921afb08'
```


Tests for q4.1 (points: 1)
```
>>> var_names = ['pop' , 'CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(pop_q4_1) == 'd0ecac9d532783c816b4995f1327a2e6'
>>> assert get_hash(CO2_q4_1) == '8f484d873ab271148f32f7d0921afb08'
```


Tests for q4.2 (points: 1)
```
>>> var_names = ['pop' , 'CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(pop) == '881602830957a3cd082095af4601bb6e'
>>> assert get_hash(CO2) == 'e88a56268abff0794d40bf0f030ed684'
```


Tests for q4.3 (points: 1)
```
>>> var_names = ['per_capita']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(per_capita) == '7b6ee17515b452692f05c2c2341c91cf'
```


Tests for q4.4 (points: 1)
```
>>> assert get_hash(q4_4) == '977c8f03a4a2c1813c6bcaaa0b831963'
```
