```python
# Initialize Otter
import otter
grader = otter.Notebook("w03Alab.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session: 03-A <br></center></h2>

<h1><center>Collections<br></center></h1>

---

```python
from resources.hashutils import *
import csv
import math
```

# Question 1: Are you available to meet?

Suppose you are writing a program to help manage your schedule. You find that you need a bit of code to check whether you are available to meet on a given day at a given time. Your calendar is available in a dictionary called `calendar`. The keys in this dictionary are your work days: `'Monday'` through `'Friday'`. The value for each key is the interval of time when you are available to begin a meeting. So, for example, `calendar['Monday'] = [8,12]` means that you are available to start a one-hour meeting on Mondays any time between 8am and 12 noon (inclusive), and not outside of those hours. 

Write code that sets `result` to `True` if you are available to meet on the day and time specified in the `day` and `time` variables. Otherwise set `result` to `False`. 

The cell below has some starter code for you to work with. Once you are satisfied with your code, you can copy-paste it into the function called `available` in the next cell.

**Note**: 
+ We are using a 24-hour clock format, so 15 corresponds to 3 p.m.
+ Your code should return `False` if the provided key is not in the `calendar` dictionary. 

**Hints**: 
+ You can check whether the provided day is a key in the dictionary with `calendar.keys()` and the membership operator `in`.
+ It is possible to do this in one line of code, but not necessary. You may define new variables to break up the work. 

```python
calendar = {
    "Monday":[8,12],
    "Tuesday":[10,12],
    "Wednesday":[13,15],
    "Thursday":[13,15], 
    "Friday":[8,15]
}

day = "Monday"
time = 12

result = ...
```

Paste your code into the next cell. Do not set the `day` and `time`, since these are passed as arguments to `available`.

```python
def available(day, time):
    ...
```

```python
# Test it
available("Monday", 7)
```

```python
grader.check("q1")
```

# Question 2: Counting votes

The next cell loads the two lists that we will work with in this question (we will learn more about loading data from csv files later in the course). The two lists: `voters` and `preferred_candidates`, respectively contain the names of the voters from a recent election and the candidates that they voted for (suppose that we are privy to such information). The two lists line up so that

`voters[i]` voted for `preferred_candidates[i]`

for every index `i`. 

```python
# RUN THIS CELL AS-IS. DO NOT MODIFY.
with open('resources/voters.csv', 'r') as file:
   voters = list(csv.reader(file))
voters = voters[0]     
with open('resources/preferred_candidates.csv', 'r') as file:
   preferred_candidates = list(csv.reader(file))
preferred_candidates = preferred_candidates[0] 
```

## Question 2.1

How many votes were collected? 

**Hint**: [len](https://docs.python.org/3/library/functions.html#len)

```python
num_votes_collected = ...
num_votes_collected
```

```python
grader.check("q2p1")
```

## Question 2.2

Who is Liv Goel's preferred candidate?

**Hint**: 
+ You should first find Liv Goel's index in the list of voters, and then use this to find their preferred candidate. 
+ You can use the [index](https://docs.python.org/3/tutorial/datastructures.html) method of Python lists.

```python
...
Liv_Goel_vote = ...
```

```python
grader.check("q2p2")
```

## Question 2.3

How many *unique* voter names were there? How many unique candidates were there?

**Hint**: 
+ You can cast your list to a Python [set](https://docs.python.org/3/tutorial/datastructures.html#sets) to obtain the collection of unique voters.

```python
num_unique_voters = ...
num_unique_candidates = ...
num_unique_voters
```

```python
grader.check("q2p3")
```

Noting that there were more votes collected than voters, it seems that there was at least one person who voted more than once. In the next lab we will figure out who that was, as well as the winner of the election.

# Question 3: List slicing

Next we will practice the technique of *slicing* for extracting sub-lists from lists. The cell below loads a list called `A`. It consists of 1000 randomly generated numbers between 0 and 1. 

```python
# RUN THIS CELL AS-IS. DO NOT MODIFY.
with open('resources/random_numbers.csv', 'r') as file:
   A = list(csv.reader(file,quoting = csv.QUOTE_NONNUMERIC))
A = A[0]
```

## Question 3.1

Find the sum of the first 20 entries of `A`. 

Assign it to the variable `sol1`.

**Hint**: 
+ Python's builtin [sum](https://docs.python.org/3/library/functions.html#sum).

```python
sol1 = ...
```

```python
grader.check("q3p1")
```

## Question 3.2

Find the maximum of the last 20 entries in `A`. 

Assign it to the variable `sol2`.

**Hint** 
+ Python's builtin [max()](https://docs.python.org/3/library/functions.html#max)

```python
sol2 = ...
```

```python
grader.check("q3p2")
```

## Question 3.3

Create a list with every 120th entry in `A`. That is, the list should include `A[0]`, then skip 119 values, then include `A[120]`, then skip 119 values, and so on. 

Assign it to `sol3`.

```python
sol3 = ...
```

```python
grader.check("q3p3")
```

## Question 3.4

Find the absolute value of the difference between the sums of the even and odd-indexed entries of `A`.

That is:

$$ \bigg| (A[0]+A[2]+...+A[998]) - (A[1]+A[3]+...+A[999])  \bigg| $$

Assign it to `sol4`.

**Hint**:
+ Python's builtin [abs()](https://docs.python.org/3/library/functions.html#abs)

```python
absdiff = ...
```

```python
grader.check("q3p4")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> available('Monday', 9.5) and (not available('Monday', 7))
True
```
```
>>> not available('Funday', 13)
True
```
```
>>> available('Friday', 10) and (not available('Wednesday', 16))
True
```
```
>>> available('Thursday', 13) and available('Thursday', 15)
True
```


Tests for q2p1 (points: None)
```
>>> num_votes_collected > 3200 and num_votes_collected < 3300
True
```
```
>>> get_hash(num_votes_collected, 4) == '0d352b4d3a317e3eae221199fdb49651'
True
```


Tests for q2p2 (points: None)
```
>>> len(Liv_Goel_vote) == 14
True
```
```
>>> get_hash(Liv_Goel_vote) == '7054f8890aa9e7895e5ff73cf3191b3e'
True
```


Tests for q2p3 (points: None)
```
>>> num_unique_voters > 3100 and num_unique_voters < 3200
True
```
```
>>> num_unique_candidates > 0 and num_unique_candidates < 100
True
```
```
>>> get_hash(num_unique_voters) == '721e7285b298cde5b3d0c973ed8d7b63' and get_hash(num_unique_candidates) == '8f14e45fceea167a5a36dedd4bea2543'
True
```


Tests for q3p1 (points: None)
```
>>> sol1 > 9 and sol1 < 10
True
```
```
>>> get_hash(sol1, 4) == 'e3a71a2703aa3ce96a7f57c534e39e4f'
True
```


Tests for q3p2 (points: None)
```
>>> sol2 > 0 and sol2 < 1
True
```
```
>>> get_hash(sol2, 4) == 'f470b40b7d4a565c1d08911e0f6bd42e'
True
```


Tests for q3p3 (points: None)
```
>>> len(sol3) == 9
True
```
```
>>> get_hash(sol3, 4) == '6963dc8970ed9023c081f891298627e4'
True
```


Tests for q3p4 (points: None)
```
>>> absdiff > 8 and absdiff < 9
True
```
```
>>> get_hash(absdiff, 4) == 'a380656244f4cdbaec8adf84a2881719'
True
```
