```python
# Initialize Otter
import otter
grader = otter.Notebook("w12Alab.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session #12-A <br></center></h2>

<h1><center>Estimating the volume of liquid in a spherical mixer<br></center></h1>

---

```python
import numpy as np
import matplotlib.pyplot as plt
from resources.hashutils import *
```

The figure below shows a spherical glass container used to mix liquids. The container has a hole through which the liquids are poured, and a small propeller to do the mixing. Unfortunately the container is old and the tick marks indicating the volume of liquid are not legible. The only way to know the amount of liquid in the container is to measure the height $h$ of the fluid's surface, and then to use the formula for the volume of a partially filled sphere:

$$ V = \frac{\pi h^2}{3} \left(3r - h \right) $$

Here $r=10$ cm is the radius of the sphere and $V$ is the volume in cubic centimeters (i.e. milliliters)

<center><img src="resources/sphere.png" width=300/></center>


## Part 1: Volume evaluation

Define a global constant `r` for the radius (in centimeters). Then  write a function called `evalV(h)` that takes the fluid level `h` (in centimeters), and returns the volume of fluid in milliliters, according to the formula just given. 

```python
...
```

```python
grader.check("p1")
```

## Part 2: Total volume

What is the total volume of the sphere in liters? Save this value as a global variable called `total_volume`.

```python
...
```

```python
grader.check("p2")
```


A chemical recipe requires that you mix 3 liters of a liquid using the glass mixer. You decide to use the root-finding methods learned in class to find the level $h$ corresponding to this volume of liquid. You begin by plotting the function $V(h)$ in order to get an approximate sense of the answer.

## Part 3: Plot volume versus height

Make the plot shown in the figure. The plot shows
+ the volume of liquid in milliliters ($V(h)$) as a function of height $h$ in centimeters (magenta line), and
+ the desired volume (dotted blue line)

Additional specifications:

+ Use `ax1` for the axes object. 
+ Use `figsize=(8,4)`
+ The computed volume of liquid in the mixer (magenta line) should be evaluated on a grid of 50 points ranging from zero to $2r$ inclusive. 
+ The grid and legend should be as shown in the plot.
+ Remove all four spines (set their visibility to `False`).
+ Plot the horizontal line using `plot` or [`axhline`](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.axhline.html), but **not hlines**, otherwise the autograder will fail.

<center><img src="resources/volvsheight.png" width=600 /></center>


```python
fig1, ax1 = plt.subplots(...)
...
```

```python
grader.check("p3")
```

## Part 4: Equation to solve

We can see in the figure above that 3 liters of fluid correspond to a height of a little more than 12.5 centimeters. To solve this more precisely, we must find the exact point at which the function $V(h)$ crosses the 3000 ml line. In other words, we must find the roots of the function $f(h)$:

$$ f(h)=V(h)-3000$$

Define the function `evalf(h)` that evaluates $f(h)$ in milliliters.

```python
...
```

```python
grader.check("p4")
```

## Part 5: Number of bisection steps

Write a function called `compute_bisection_steps(N0,P0,epsilon)` that computes the number of bisection steps needed to reduce the interval width from `|P0-N0|` to less or equal than a given `epsilon`. The formula for doing this was derived in lecture and is reproduced below.

$$ k = \text{ceil} \left( \log_2 \left( \frac{ |P_0-N_0| }{ \epsilon} \right)  \right) $$

**Hint**:
+ NumPy's [ceil](https://numpy.org/doc/stable/reference/generated/numpy.ceil.html) method returns a float. You should cast this to an integer in order to avoid problems in later parts.
+ Return `0` whenever $|P_0-N_0|<\epsilon$. 

```python
...
```

```python
# You can use this cell to test your code

```

```python
grader.check("p5")
```

## Part 6: Bisection

Implement the bisection algorithm according to the description given in lecture. Below we provide a template for your code with comments. 

The `bisection` method takes as inputs the initial interval values `N0` and `P0`, and the desired precision `epsilon`. Here, `N0` is a value of $h$ corresponding to a *negative* evaluation of $f(h)$, and `P0` is a value of $h$ corresponding to a *positive* evaluation of $f(h)$. Your `bisection` function should return two NumPy arrays corresponding to the *negative* and *positive* sequences (in that order).

**Note**
+ In the case where $|P_0-N_0|<\epsilon$, the code should return `(np.array(N0), np.array(P0))`.
+ You can use functions already defined in previous parts for certain steps of the algorithm. 

```python
def bisection(N0,P0,maxwidth):

    # Compute the number of steps to execute.
    numsteps = ...

    # Deal with the case that no bisection steps are required.
    ...

    # Allocate the negative and positive sequences using np.empty.
    # These sequences should have length = numsteps+1.
    ...

    # Initialize the zeroth position of the negative and positive sequences.
    ...

    # Iterate with k ranging from 1 to numsteps (inclusive)
    for k in ...:

        # Find midpoint of the negative and positive sequences.
        ...
        
        # evaluate f at the midpoint.
        ...

        # compute the next k'th value for the negative and positive sequences.
        ...
        ...

    # return the results
    ...
```

```python
# You can use this cell to test your code

```

```python
grader.check("p6")
```

## Part 7: Plot

Plot the two sequences computed by the bisection method with parameters `N0=0`, `P0=20`, and `epsilon=0.1`.
Define the solution `hsol` as the average of the last values in the two sequences. That is, `hsol=0.5*(N[-1]+P[-1])`, where `N` and `P` are the two arrays returned by `bisection`.

Plot specifications:

+ Use `ax2` for the axes object. 
+ Label the three lines as shown in the figure below. 
+ The `N` and `P` lines have `linewidth=2` and `marker='o'`.  
+ Draw a horizontal line at the level of `hsol`. Again, use `axhline` or `plot`, but not `hlines`.

<center><img src="resources/bisection.png" width=600/></center>

```python
...
```

```python
grader.check("p7")
```

## Part 8: Derivative of $f(h)$


Next we will apply the Newton-Raphson algorithm to the same problem. This is possible since the function that we are solving is differentiable, and the derivative can be easily evaluated in code. Implement the function `evalfprime` that returns $f'(h)$: the derivative of $f$ for a given value of $h$. 

Recall:

\begin{align*} 
&V(h) = \frac{\pi h^2}{3} \left(3r - h \right) \\
&f(h)=V(h)-3000
\end{align*}



```python
...
```

```python
grader.check("p8")
```

## Part 9: Newton-Raphson

Implement the Newton-Raphson algorithm as described in lecture.

Write a function `newton_raphson(x0,epsilonx,epsilonf,maxsteps)` that takes an initial guess `x0`, tolerances `epsilonx` and `epsilonf`, and the maximum number of steps `maxsteps`.The function should return a NumPy array with the resulting sequence of values, beginning with `x0`.

**Note**. 

1. You will notice that the steps do not exactly follow the pseudo-code given in class. The pseudo-code involves a "repeat-until", which in Python would be implemented with a "while" loop. However we do it here with a "for" loop. The reason for this is that when using NumPy it is better to pre-allocate the entire array (e.g. with np.empty) than to extend the array in each iteration. But once we've allocated a fixed-size array, it is better to iterate through it with a "for" loop instead of a "while" loop.

2. Note also that the function is cubic and therefore may have up to three real roots. In fact, this function *does* have three real roots, but only one of them is in the acceptable range of heights, between 0 and 20 cm. This is ok. You should allow the Newton-Raphson method to return these unphysical solutions if it finds them.

**Hint**. Here are a few NumPy functions that may be useful for removing the trailing NaN's in `x`.
+ [`np.isnan`](https://numpy.org/doc/2.1/reference/generated/numpy.isnan.html) Creates a boolean mask with `True`'s marking the locations of NaN's.
+ [`np.where`](https://numpy.org/doc/2.2/reference/generated/numpy.where.html) Returns the indices from a boolean mask corresponding to `True` entries.
+ [`np.min`](https://numpy.org/doc/2.2/reference/generated/numpy.min.html) Returns the smallest value in an array.

```python
def newton_raphson(x0,epsilonx,epsilonf,maxsteps):

    # Allocate x by creating an array of np.nan values of length maxsteps
    ...

    # Put the initial guess in the first position of x
    ...

    # Iterate through x, starting at position 1
    for k in ...:

        # Find the next value in the sequence according to the Newton-Raphson update formula
        ...
                
        # Exit the loop if the difference between the latest two values in the sequence is less than epsilonx
        ...

        # Exit the loop if the function evaluation is within epsilonf of zero.
        ...

    # Remove any trailing NaNs (remaining if the loop was exited before reaching maxsteps).
    ...

    return x
```

```python
# Use this cell to test your code
```

```python
grader.check("p9")
```

## Part 10: Compare Bisection and Newton-Raphson

We can now test both methods to see which converges in fewer steps. Note that this does not imply that it takes less *time*, since the steps of one algorithm may be more computationally expensive (and therefore *slower*) than the steps of another algorithm. A Newton-Raphson step is more expensive than a bisection step, since it involves evaluations of both $f(h)$ and $f'(h)$, instead of just $f(h)$.

Make the plot shown below that overlays the progression of both algorithms. Here is the specification. 

+ Use `ax3` for the axes object. 

For bisection:
+ Run the algorithm with `N0=0`, `P0=20`, and `epsilon=0.1`.
+ Plot the upper and lower limits of the interval with solid lines and 'o' markers. 
+ Use the [ax.fill_between](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.fill_between.html) method to shade in the space between the two limits (I like to use a transparency value of `alpha=0.2` for this). 
+ Label the "fill_between" element (not the upper and lower lines) as "Bisection".

For Newton-Raphson:
+ Run the algorithm with `x0=5`, `epsilonx=0.1`, `epsilonf=0`, `maxsteps=10`.
+ Plot the sequence with a solid line and 'o' markers. 
+ Label it as "Newton-Raphson".

For the solution.
+ Consider the last value of the Newton-Raphson sequence to be the solution. 
+ Plot a horizontal dashed line that goes through the solution. 
+ Label it as "solution".

<center><img src="resources/compare.png" width=600/></center>

```python
...
```

```python
grader.check("p10")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for p1 (points: None)
```
>>> np.isclose(evalV(1), 30.36872, atol=0.01)
np.True_
```
```
>>> get_hash([evalV(h) for h in np.arange(1, 6)], 5) == 'cab0173151fc8d51e6d2fb7b2c16962f'
True
```


Tests for p10 (points: None)
```
>>> z = ax3.get_legend_handles_labels()
>>> fb = z[0][z[1].index('Bisection')]
>>> path = fb.get_paths()[0]
>>> get_hash(np.sort(np.unique(path.vertices[:, 1])), 4) == '4c6ec1d3a33c0af61838e7c4f5b88224'
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax3.get_lines()}
>>> get_hash(plotdata['solution'][0], 4) == 'e5f53791e854c3b33cf13e30406312a2'
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax3.get_lines()}
>>> get_hash(plotdata['Newton-Raphson'], 4) == 'a619bcae08752f158de2e2d07ea47d24'
True
```
```
>>> {line.get_marker() for line in ax3.get_lines() if line.get_label() != 'solution'} == {'o'} and {line.get_linestyle() for line in ax3.get_lines() if line.get_label() != 'solution'} == {'-'}
True
```


Tests for p2 (points: None)
```
>>> 'total_volume' in globals()
True
```
```
>>> total_volume > 0 and total_volume < 10
True
```
```
>>> get_hash(total_volume, 4) == 'fd4c55c738a8e8c110920d67f21e93b9'
True
```


Tests for p3 (points: None)
```
>>> import matplotlib
>>> 'ax1' in globals() and isinstance(ax1, matplotlib.axes._axes.Axes)
True
```
```
>>> np.all(ax1.get_figure().get_size_inches() == [8.0, 4.0])
np.True_
```
```
>>> {line.get_label().lower() for line in ax1.get_lines()} == {'desired volume', 'volume of liquid in the mixer'}
True
```
```
>>> z = [line for line in ax1.get_lines() if line.get_label().lower() == 'volume of liquid in the mixer']
>>> np.all(np.isclose(z[0].get_xdata(), np.linspace(0.0, 20.0, 50), atol=0.0001))
np.True_
```
```
>>> z = [line for line in ax1.get_lines() if line.get_label().lower() == 'volume of liquid in the mixer']
>>> get_hash(z[0].get_ydata(), 4) == '64e90d44e0bc19fbd9e4a00b0fea093e'
True
```
```
>>> line2 = list(filter(lambda x: x.get_label() == 'desired volume', ax1.get_lines()))[0]
>>> ydata = np.array(line2.get_ydata())
>>> np.all(ydata == 3000) and len(ydata) > 1
True
```
```
>>> np.all([l.get_data()[1][0] for l in ax1.get_ygridlines()] == np.arange(-1000, 5500, 1000))
np.True_
```
```
>>> np.all([l.get_data()[0][0] for l in ax1.get_xgridlines()] == np.arange(-2.5, 23, 2.5))
np.True_
```
```
>>> ax1.get_legend().get_visible()
True
```


Tests for p4 (points: None)
```
>>> evalf(0) == -3000 and evalf(20) > 1000
True
```
```
>>> get_hash([evalf(h) for h in range(10, 16)], 4) == '94436faad921ede2ea6ec58b2e554f4b'
True
```


Tests for p5 (points: None)
```
>>> type(compute_bisection_steps(0, 20, 1)) == int
True
```
```
>>> compute_bisection_steps(0, 0.4, 1) == 0
True
```
```
>>> [compute_bisection_steps(0, p, 1) for p in np.arange(20, 60, 4)] == [5, 5, 5, 5, 6, 6, 6, 6, 6, 6]
True
```
```
>>> np.all([compute_bisection_steps(0, p, 1) for p in np.arange(20, 100, 7)] == [5, 5, 6, 6, 6, 6, 6, 7, 7, 7, 7, 7])
np.True_
```


Tests for p6 (points: None)
```
>>> N, P = bisection(0, 10, 3)
>>> np.all(N == np.array([0.0, 5.0, 7.5])) and np.all(P == np.array([10.0, 10.0, 10.0]))
np.True_
```
```
>>> N, P = bisection(0, 0.3, 0.5)
>>> N == np.array(0) and P == np.array(0.3)
np.True_
```
```
>>> N, P = bisection(0, 20, 0.7)
>>> get_hash(N, 3) == 'ea05816119f2e8bbc0ab6de93e6009d7' and get_hash(P, 3) == 'f7fd87545114eae97bb3695b349a6cea'
True
```


Tests for p7 (points: None)
```
>>> 'ax2' in globals()
True
```
```
>>> len(ax2.get_lines()) == 3
True
```
```
>>> {line.get_label() for line in ax2.get_lines()} == {'N', 'P', 'hsol'}
True
```
```
>>> xdata = [line.get_xdata() for line in ax2.get_lines() if line.get_label() == 'N'][0]
>>> np.all(xdata == np.arange(9))
np.True_
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> get_hash(plotdata['hsol'][0], 4) == '1a731a46d22a4030cab90d4195426573'
True
```
```
>>> plotdata = {line.get_label(): line.get_ydata() for line in ax2.get_lines()}
>>> get_hash(plotdata['N'][0], 4) == '30565a8911a6bb487e3745c0ea3c8224' and get_hash(plotdata['P'][0], 4) == '75cf3ac5e70c76583be3efb5012bd44e'
True
```
```
>>> plotdata = {line.get_label(): line for line in ax2.get_lines()}
>>> plotdata['N'].get_linewidth() == 2.0 and plotdata['N'].get_marker() == 'o' and (plotdata['P'].get_linewidth() == 2.0) and (plotdata['P'].get_marker() == 'o')
True
```


Tests for p8 (points: None)
```
>>> evalfprime(10) > 300 and evalfprime(10) < 320
True
```
```
>>> get_hash([evalfprime(h) for h in np.linspace(10, 20, 5)], 5) == '9e45fa83f5b94d7392168566c0fd34cb'
True
```


Tests for p9 (points: None)
```
>>> x = newton_raphson(x0=10, epsilonx=1, epsilonf=1, maxsteps=10)
>>> x[0] == 10
np.True_
```
```
>>> x = newton_raphson(x0=10, epsilonx=1, epsilonf=1, maxsteps=10)
>>> len(x) == 3
True
```
```
>>> x = newton_raphson(x0=1, epsilonx=1, epsilonf=1, maxsteps=5)
>>> len(x) == 5
True
```
```
>>> x = newton_raphson(x0=1, epsilonx=0.1, epsilonf=0.1, maxsteps=100)
>>> get_hash(x, 4) == 'cc67c66000938d314e4801d7ab53358e'
True
```
```
>>> x = newton_raphson(x0=-10, epsilonx=0.1, epsilonf=1, maxsteps=10)
>>> get_hash(x, 4) == 'dcc0ecf5a677a9fe6a8bfbe0332ec63f'
True
```
