```python
# Initialize Otter
import otter
grader = otter.Notebook("w01disc.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Discussion session #01 <br></center></h2>

---


```python
from resources.hashutils import *
import math
```

# Naviagating Jupyter notebooks

At any time, the notebook can either be in <font color=green>**edit mode**</font> or in <font color=dodgerblue>**command mode**</font>. You can toggle between these two by pressing <kbd>Enter</kbd>  for <font color=green>**edit mode**</font> and <kbd>Esc</kbd> for <font color=dodgerblue>**command mode**</font>.

<font color=green>**Edit mode**</font> is for editing the contents of a cell, while <font color=dodgerblue>**command mode**</font> is for cell-level manipulations, such as creating, deleting, and rearranging cells. 

Here is a list of useful keyboard shortcuts to use in <font color=dodgerblue>**command mode**</font>:
* <kbd>A</kbd>: Insert a new cell above
* <kbd>B</kbd>: Insert a new cell below
* <kbd>X</kbd>: Cut the cell
* <kbd>C</kbd>: Copy the cell
* <kbd>V</kbd>: Paste the copied cell
* <kbd>Y</kbd>: Change a cell to a code cell
* <kbd>M</kbd>: Change a cell to a markdown cell

These actions are also available as toolbar buttons and dropdown menus.

# Executing cells

+ Both code and markdown cells can be executed with keyboard shortcuts.

    + <kbd>Shift</kbd>+<kbd>Enter</kbd> runs the currently highlighted cell(s), and then advances to the next cell. The `▶ Run` button in the toolbar does the same.
    + <kbd>Ctrl</kbd>+<kbd>Enter</kbd> (<kbd>Cmd</kbd>+<kbd>Enter</kbd> on macOS) runs the currently highlighted cells without advancing.

+ Cells can also be executed using the **Run** menu. 

+ It is good to get in that habit of restarting your kernel when running all cells. Do this with **Kernel**->**Restart Kernel and Run All Cells...**. 

```python
import random
print("This is a random number: ", random.random())
```

# Autograding and Gradescope

We will be using an automatic grader for lab assignments. Each question in a lab assignment will have one or more **solution cells** followed by a **test cell**. You will add your code to the **solution cells** as instructed in the lab, and then test your answer by running the **test cell**. The **test cell** hides a series of tests, and produces output indicating whether your solution is correct. 


## Example 1.
Given two variables, `a=1` and `b=2`, compute their sum and assign it to variable `c`. 

+ The correct solution to this exercise is to replace the `...` below with `a+b`.  Try that out, and see that the testing cell marks your answer as correct. 
+ Also try an *incorrect* answer, by replacing the `...` with `-2`. Notice what happens when you run the testing cell.

#### Graded cell:

```python
a = 1
b = 2
c = ...
```

#### Testing cell:

```python
grader.check("q1")
```

# Example 2. Functions

One of the most important concepts in programming is that of the *function*. A function is a way of packaging pieces of code for reuse. 

Instruction: Write a function called `area_of_rectangle` that takes as input the lengths `a` and `b` of a rectangle and returns the area of the rectangle.

```python
...
```

```python
grader.check("q2")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> c > 0
True
```
```
>>> get_hash(c, 4) == '55c82b601deae028c1c5e87fd820923d'
True
```


Tests for q2 (points: None)
```
>>> math.isclose(area_of_rectangle(1, 2), 2.0)
True
```
```
>>> get_hash(area_of_rectangle(4.54, 2.43), 4) == '68ffa29331cc2e99c83dd6567daf04cd'
True
```
