```python
# Initialize Otter
import otter
grader = otter.Notebook("w04Blab.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session: Week 4-B <br></center></h2>

<h1><center>Operations with Numpy arrays<br></center></h1>

---

```python
from resources.hashutils import *
import numpy as np
import matplotlib.pyplot as plt
```

# Part 1: 1D heat conduction

Here we will analyze some data that was collected from three temperature sensors placed along the length of a metal rod. The sensors are labeled s0, s1, and s2 in the figure below. A flame was placed under sensor s0, causing it to heat up to about $600^{\circ} C$ (degrees celsius, or centigrade). At the opposite end, a block of ice keeps sensor s2 at around $0^{\circ}C$.

<center><img src="resources/rod_diagram.png" width="300" /></center>

The plots below shows measurements recorded at several time instances from the three sensors. We can see that s0 and s2 remain approximately fixed at $600^{\circ} C$ and  $0^{\circ} C$, while measurements from s1 gradually rise from $0^{\circ} C$ to about $300^{\circ} C$.

<center><img src="resources/rod_diff2Dlabeled.png" width="700" /></center>

Run the next cell to load this data into a NumPy array called `T`. 

`T` has one column for each sensor, and one row for each instant of time.

```python
T = np.loadtxt('resources/1drod.csv', delimiter=',')
T
```

## Part 1.1 -- Number of time instants

How many time instants were recorded? Save this value to the variable `N`. 


**Note**. Don't simply count the rows by hand. Use the [`shape`](https://numpy.org/doc/stable/reference/generated/numpy.shape.html) method of the NumPy array. 


```python
N = ...
```

```python
grader.check("p1p1")
```

## Part 1.2 -- Temperature in degrees fahrenheit

Convert the temperature measurements to degrees Fahrenheit. Recall that the formula for the conversion is:

$$ T_f = \left(9/5\right) T + 32 $$

where $T$ is a temperature value in celsius, and $T_f$ is the corresponding value in fahrenheit. Store the resulting $N\times 3$ array as `Tf`.

```python
Tf = ...
```

```python
grader.check("p1p2")
```

## Part 1.3 -- Average over time

What is the average of the temperature readings for each of the three sensors in degrees Celsius? Store this as `Tmean0`. It should be an array with 3 entries (one per sensor). 

```python
Tmean0 = ...
```

```python
grader.check("p1p3")
```

## Part 1.4 -- Spatial average

The spatial average temperature is the mean of the three sensor readings at a particular time. Compute this value in degrees Celsius for each time instant and store it as `Tmean1`. This should be an array with `N` entries (one per time instant). 

```python
Tmean1 = ...
```

```python
grader.check("p1p4")
```

## Part 1.5 -- Standard deviations

Sensors s0 and s2 are supposedly kept at constant temperatures, but we can see from the figure that there is some fluctuation or *noise* in the measurements. We can quantify this noise in the signal by computing its standard deviation. The standard deviation $\sigma$ of an array $A=[a_1,...,a_N]$ is defined as:

$$ \sigma = \sqrt{ \frac{1}{N-1} \sum_{i=1}^{N} (a_i-\bar{a})^2 } $$

Here $\bar{a}$ is the mean of the elements in $A$, and $N$ is the number of elements in $A$.

Compute the standard deviations for each of the sensors, and store the result as `Tstd`. 

**Notes**:
+ The shape of `Tstd` should be `(3,)`.
+ Try to do this in "the NumPy way", using broadcasting and methods such as [np.sum()](https://numpy.org/doc/stable/reference/generated/numpy.sum.html) instead of loops. 
+ *This formula can be implemented in one line of code*.
+ **Do not use [np.std()](https://numpy.org/doc/stable/reference/generated/numpy.std.html)**

```python
Tstd = ...
```

```python
grader.check("p1p5")
```

# Part 2: Bouncing ball

There are physical systems, such as the vibrating rod of lab 3A, whose evolution through time is simple enough that it can be expressed with a single formula. Other more complicated dynamics are more difficult to write down. A ball moving in a vacuum that is subject to gravity follows a trajectory that is easily described as a parabola. However, if we introduce the element of a "floor", and a model for "bouncing" and "rolling", then the picture becomes more complicated. It is no longer possible to compute the trajectory of the ball with a single line of code. We must use a `for` loop.

Here we will use a standard coding pattern for computing the trajectory of a dynamical system through time. This evolution is shown in the figure below. The ball follows a sequence of parabolic arcs, separated by bounces (red dots) where some of its energy is lost. Eventually the ball losing most of its energy and starts rolling (green line).

<img src="resources/ball_pxpy.png" width="600" />


## Part 2.1 -- Time instants

Create an array of time instants called `time` ranging from 0 to 2 seconds and with a time step of 0.015 seconds. Find the number of elements in `time` and call it `K`.

**Notes**: 
+ Use [`np.arange`](https://numpy.org/doc/stable/reference/generated/numpy.arange.html).
+ `time` will not include a 2, since 2 is not a multiple of `dt`. Rather, the last value in `t` should be 1.995.

```python
time = ...
K = ...
```

```python
grader.check("p2p1")
```

## Part 2.2 -- Memory allocation


The position and speed of the ball at any given time step `k` is described with four numbers:

+ `px[k]` ... the horizontal component of its position (in meters)
+ `py[k]` ... the vertical component of its position (in meters)
+ `vx[k]` ... the horizontal component of its speed (in meters per second)
+ `vy[k]` ... the vertical component of its speed (in meters per second)


We will also want to track the bounces of the ball. To do this we will use a boolean array called `bounce` that is of the same length as `time` and is `True` when a bounce occurs (and `False` elsewhere). The plot above shows bounce times as red dots.

Because a ball cannot bounce while it is rolling, we will also want to track the rolling state of the ball. This will be captured with a boolean variable `rolling[k]`, which is `True` if the ball is rolling, and `False` otherwise. 

Next we will allocate memory for these six arrays: `px`, `py`, `vx`, `vy`, `inflight`, and `bounce`. 

**Notes**. 
+ Useful methods for this include [np.empty](https://numpy.org/doc/stable/reference/generated/numpy.empty.html ) and [np.full](https://numpy.org/doc/stable/reference/generated/numpy.full.html).
+ Make sure that your arrays are of the correct data type. The position and velocity arrays contain floats, while `rolling` and `bounce` arrays contain booleans. You can check this with e.g. `rolling.dtype`.


```python
px = ...
py = ...
vx = ...
vy = ...
rolling = ...
bounce = ...
```

```python
grader.check("p2p2")
```

## Part 2.3 -- State initialization

The ball begins its flight (time step `k=0`) at a horizontal position of 0 and a vertical position of 2 meters. Its initial horizontal speed is 1 m/s and its vertical speed is 3 m/s. This means that initially it is not rolling and not bouncing. Record this initial state information into the six state arrays.

```python
...
```

```python
grader.check("p2p3")
```

## Part 2.4 -- Loop through time

We can now code the "for" loop that will compute the trajectory of the ball through time. In the `k`'th iteration, the task is to compute values to write into the `k`'th position of our six arrays, based on the values in the `k-1`'th positions.

The calculation is done in four steps:

1. **Rolling update**. The ball is considered to be rolling at time instant `k` if its position at `k-1` is within one centimeter of the floor *and* its vertical speed at `k-1` is less than 0.01 m/s in absolute value. 

2. **Position update**. We assume that `vx[k-1]` and `vy[k-1]` are the *average* velocity of the ball over the time interval from `time[k-1]` to `time[k]`. The ball therefore advances by `dt*vx[k-1]` in that interval, and hence `px[k] = px[k-1]+dt*vx[k-1]` (here `dt` is the time step of 0.015 second). The same is true for the vertical position `py[k]`, provided the ball is not rolling at time `k-1`. If it is rolling, then the vertical position at time `py[k]` is set to zero.

2. **Velocity update**. The horizontal speed of the ball remains constant for all time. The ball's change in vertical speed from `k-1` to `k` depends on whether it was rolling or not at time `k-1`. If it was rolling, then its vertical speed is zero. Otherwise its vertical speed is decreased by `dt*g` with respect to the previous time instant, where `g` is the acceleration of gravity, ie. $9.81 \; m/s^2$.


4. **Bounce detection**. A bounce occurs at time `k` whenever all three of the following conditions are met:
    + The vertical velocity at time `k` is less-or-equal than zero,
    + The vertical position at time `k` is less-or-equal than zero,
    + The ball is not rolling at time `k`.

5. **Bounce action**. If a bounce is detected at time `k`, then 
    + The vertical velocity is flipped from negative to positive, and multiplied by an energy-dissipating factor of 0.3.
    + The vertical position at time `k` should be set to 0.

```python
dt = 0.015

for k in range(1,K):

    # 1. rolling update
    rolling[k] = ...

    # 2. position update
    px[k] = ...
    py[k] = ...   # feel free to use an `if` statement here

    # 3. velocity update
    vx[k] = ...
    vy[k] = ...

    # 4. bounce detection
    bounce[k] = ...

    # 5. bounce action
    if bounce[k]:
        ...
```

Use this cell to plot your result. Your plot should look identical to the one in the problem description

```python
fig, ax = plt.subplots(figsize=(8,3))
plt.axhline(0,c='k',linewidth=0.5)
plt.plot(px, py, 'k-')
plt.plot(px[rolling], py[rolling], 'g',linewidth=4,label='rolling')
plt.plot(px[bounce], py[bounce], 'ro',label='bounce')
ax.spines[['top','right']].set_visible(False)
plt.ylim(-.4,2.5)
plt.xlim(0,2)
plt.legend(fontsize=14)
plt.xlabel('px',fontsize=14)
plt.ylabel('py',fontsize=14)
```

```python
grader.check("p2p4")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for p1p1 (points: None)
```
>>> N > 15 and N < 30
True
```
```
>>> get_hash(N, 2) == '6bc071ec71e51c704acd13cdc898fd93'
True
```


Tests for p1p2 (points: None)
```
>>> Tf.shape == (23, 3) and np.all(Tf > -10) and np.all(Tf < 1200)
np.True_
```
```
>>> get_hash(Tf, 4) == '52850960631167ab9dc86c9538e055ab'
True
```


Tests for p1p3 (points: None)
```
>>> Tmean0.shape == (3,) and np.all(Tmean0 > 0)
np.True_
```
```
>>> get_hash(Tmean0, 3) == 'c616f7c34a1ffebe45f40fa7ec05b73e'
True
```


Tests for p1p4 (points: None)
```
>>> Tmean1.shape == (N,) and np.all(Tmean1 > 0)
np.True_
```
```
>>> get_hash(Tmean1, 4) == 'c7997717c0e9d48f6ec1970728c17772'
True
```


Tests for p1p5 (points: None)
```
>>> Tstd.shape == (3,) and np.all(Tstd > 0)
np.True_
```
```
>>> get_hash(Tstd, 3) == '7064fcdd970094099b562e0a6aa15ab8'
True
```


Tests for p2p1 (points: None)
```
>>> np.all(time >= 0) and np.all(time <= 2) and (K > 0)
True
```
```
>>> get_hash(time, 3) == '7d731d2cf5aaed612a09b3249c2d5f24' and get_hash(K, 3) == '7a8e1d9a305be336af766f37377c2639'
True
```


Tests for p2p2 (points: None)
```
>>> px.dtype == float and py.dtype == float and (vx.dtype == float) and (vy.dtype == float)
True
```
```
>>> rolling.dtype == bool and bounce.dtype == bool
True
```
```
>>> get_hash(px.shape, 3) == '01b18d0cddc781117c77bfddeefe0d1f' and get_hash(py.shape, 3) == '01b18d0cddc781117c77bfddeefe0d1f'
True
```
```
>>> get_hash(vx.shape, 3) == '01b18d0cddc781117c77bfddeefe0d1f' and get_hash(vy.shape, 3) == '01b18d0cddc781117c77bfddeefe0d1f'
True
```
```
>>> get_hash(rolling.shape, 3) == '01b18d0cddc781117c77bfddeefe0d1f' and get_hash(bounce.shape, 3) == '01b18d0cddc781117c77bfddeefe0d1f'
True
```


Tests for p2p3 (points: None)
```
>>> get_hash(px[0], 1) == '30565a8911a6bb487e3745c0ea3c8224' and get_hash(py[0], 1) == 'd1bd83a33f1a841ab7fda32449746cc4'
True
```
```
>>> get_hash(vx[0], 1) == 'e4c2e8edac362acab7123654b9e73432' and get_hash(vy[0], 1) == '55c82b601deae028c1c5e87fd820923d'
True
```
```
>>> get_hash(rolling[0]) == '30565a8911a6bb487e3745c0ea3c8224' and get_hash(bounce[0], 1) == '30565a8911a6bb487e3745c0ea3c8224'
True
```


Tests for p2p4 (points: None)
```
>>> sum(bounce) == 4 and sum(rolling) > 10 and (sum(rolling) < 20)
np.True_
```
```
>>> get_hash(np.where(rolling), 3) == '08795e875ce6e01b1500302618deeca5'
True
```
```
>>> get_hash(np.where(bounce), 3) == 'a95bfccff9184e6923c978b3622025ec'
True
```
```
>>> get_hash(py[::20], 3) == 'd78f8bff4e8ae0ad06d7b8e5ccfa304c'
True
```
```
>>> get_hash(px[::20], 2) == 'daea30ca6ebeb50e5fd55dd3440111de'
True
```
