```python
# Initialize Otter
import otter
grader = otter.Notebook("w05Blab.ipynb")
```

```python
from resources.hashutils import *
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session: Week 5-B <br></center></h2>

<h1><center>Strings<br></center></h1>

---

# Part 1: Bank statement

In this problem we will create a bank statement from a list of transactions. The list is provided to us in the form of a string:

```python
flat_string='1,5,2024,ATM Withdrawal,w,50;1,10,2024,paycheck,d,2500;1,15,2024,groceries,w,93.85;1,20,2024,clothes,w,75;1,25,2024,entertainment,w,25.50'
```

The next few questions will guide you through the process of transforming the flat string into a nicely formatted bank statement. 

## Part 1.1: Split the string at semicolons (';') to obtain a list of transactions.

```python
transactions = ...
transactions
```

```python
grader.check("p1p1")
```

## Part 1.2: Split the first transaction into its parts.

Again use the string's [`split`](https://docs.python.org/3.3/library/stdtypes.html#str.split) method to split the first entry in `transactions` into 6 separate strings. Call these `month_str`, `day_str`, `year_str`, `itemtype_str`, `w_or_d_str`, and `amount_str`.

```python
...
```

```python
grader.check("p1p2")
```

## Part 1.3: Convert the strings into the desired data types

Use these variable names:

+ `day` is an integer
+ `month` is an integer
+ `year` is an integer
+ `is_withdrawal` is a boolean that is `True` if the transaction is a withdrawal (`w_or_d_str` equals `w`) and `False` if the transaction is a deposit (`w_or_d_str` equals `d`). 
+ `amount` is a float


```python
day = ...
month = ...
year = ...
is_withdrawal = ...
amount = ...
```

```python
grader.check("p1p3")
```

## Part 1.4: Process all transactions. Track the balance.

Write a function that applies steps 1.1 through 1.3 to each of the transactions in the list. This function takes two inputs: the raw transactions string (`flat_string` in part 1.1), and an initial balance. The output of this function is a list of tuples. Each tuple gathers information for a particular transaction. The tuple consists of:
```python
(day,month,year,is_withdrawal,amount,balance)
```

```python
def build_transactions_list(flat_string,initial_balance):

    # initialize the balance
    balance = ...

    # split the input string
    transactions = ...

    # create an empty list
    transactions_list = ...

    for transaction in transactions:

        # split the entry
        ...

        # type conversions
        day = ...
        month = ...
        year = ...
        is_withdrawal = ...
        amount = ...
        
        # update the balance
        ...

        # append this information to transactions_list
        transactions_list ...

    return transactions_list
```

```python
grader.check("p1p4")
```

## Part 1.5: Format the bank statement. 

The final task is to create a nicely formatted statement to present to the customer. Your job is to write a function that transforms a list of transactions built with the `build_transactions_list` function into a list of strings as shown below. 

<center><img src="resources/bank_statement.png" width="550" /></center>

The statment has these lines:

1. The name line has the string "Name:" followed by a space and the name of the customer. 
2. An empty string.
3. A separator line consisting of 53 dashes `-`.
4. A header. The parts of the header are:
    + The word `Date`, right justified in a string of width 12.
    + The word `Deposit`, right justified in a string of width 12.
    + The word `Withdrawal`, right justified in a string of width 12.
    + The word `Balance`, right justified in a string of width 12.
    + Each of these are separated by a `|` character.
5. Another separator line.
6. The list of transactions. Notice the formatting of each column. Notice also that the formatting is different for deposits and withdrawals. 
7. A final separator line.

You can check your solution against the image using the `pprint` function in the cell following the next one. 

```python
def format_statement(customer_name,transactions_list):

    WIDTH = 53 # This is the total width of the statement. 
               # and the number of dashes in the separator line.

    # Initialize and empty list of lines
    lines = ...

    # line 1: customer name
    lines.append(...)

    # line 2: empty
    lines.append(...)

    # line 3: separator (Hint: multiply a string by an integer)
    lines.append(...)

    # line 4: header
    lines.append(...)

    # line 5: another separator
    lines.append(...)

    # transaction lines
    for transaction in transactions_list:

        # unpack the transaction
        ...

        # format the date string
        datestr = ...

        # create a transaction line depending on withdrawal or deposit
        ...

        # append it to lines
        lines.append(...)

    # final separator
    lines.append(...)

    return lines
```

```python
from pprint import pprint

customer_name = 'Ismat Vandale'
transactions_list = build_transactions_list(flat_string,1000)
lines = format_statement(customer_name,transactions_list)
pprint(lines)
```

```python
grader.check("p1p5")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for p1p1 (points: None)
```
>>> isinstance(transactions, list) and len(transactions) == 5
True
```
```
>>> get_hash(transactions[0]) == '4011a5fbceb3001e334b4f3ea142e590' and get_hash(transactions[1]) == 'c88ca52d57ff7ef3f932c16fb9d3ce9f'
True
```


Tests for p1p2 (points: None)
```
>>> isinstance(month_str, str) and isinstance(day_str, str) and isinstance(year_str, str) and isinstance(itemtype_str, str) and isinstance(w_or_d_str, str) and isinstance(amount_str, str)
True
```
```
>>> get_hash(month_str) == 'c4ca4238a0b923820dcc509a6f75849b' and get_hash(day_str) == 'e4da3b7fbbce2345d7772b0674a318d5' and (get_hash(year_str) == '07811dc6c422334ce36a09ff5cd6fe71') and (get_hash(itemtype_str) == '49defd0f3fdc923db244fb898145e311') and (get_hash(w_or_d_str) == 'f1290186a5d0b1ceab27f4e77c0c5d68') and (get_hash(amount_str) == 'c0c7c76d30bd3dcaefc96f40275bdc0a')
True
```


Tests for p1p3 (points: None)
```
>>> isinstance(day, int) and isinstance(month, int) and isinstance(year, int) and isinstance(is_withdrawal, bool) and isinstance(amount, float)
True
```
```
>>> get_hash(day, 1) == '336669dbe720233ed5577ddf81b653d3' and get_hash(month, 1) == 'e4c2e8edac362acab7123654b9e73432' and (get_hash(year, 4) == '48325466001978e098c8d0cb216d7cef')
True
```
```
>>> get_hash(is_withdrawal) == 'e4c2e8edac362acab7123654b9e73432'
True
```
```
>>> isinstance(amount, float) and np.isclose(amount, 50.0)
np.True_
```


Tests for p1p4 (points: None)
```
>>> A = build_transactions_list(flat_string, 1000)
>>> np.all(np.isclose(A[0], (5, 1, 2024, True, 50.0, 950.0))) and np.all(np.isclose(A[1], (10, 1, 2024, False, 2500, 3450.0))) and np.all(np.isclose(A[2], (15, 1, 2024, True, 93.85, 3356.15))) and np.all(np.isclose(A[3], (20, 1, 2024, True, 75, 3281.15))) and np.all(np.isclose(A[4], (25, 1, 2024, True, 25.5, 3255.65)))
np.True_
```
```
>>> flat_string = '2,9,2022,ATM Withdrawal,w,50;2,13,2024,paycheck,d,540.45;3,19,2024,groceries,w,54.82;5,03,2024,clothes,w,75.45;7,25,2024,entertainment,w,125.50'
>>> get_hash(signif(build_transactions_list(flat_string, 1000), 6)) == '3b641790d00bbbbd9198532255073756'
True
```


Tests for p1p5 (points: None)
```
>>> lines = format_statement('Ismat Vandale', build_transactions_list(flat_string, 1000))
>>> len(lines) == 11
True
```
```
>>> lines = format_statement('Ismat Vandale', build_transactions_list(flat_string, 1000))
>>> [len(line) for line in lines] == [19, 0, 53, 53, 53, 53, 53, 53, 53, 53, 53]
True
```
```
>>> lines = format_statement('Ismat Vandale', build_transactions_list(flat_string, 1000))
>>> lines[0] == 'Name: Ismat Vandale'
True
```
```
>>> stmt = format_statement('ONWRBOIWER', build_transactions_list('2,9,2022,ATM Withdrawal,w,50;2,13,2024,paycheck,d,540.45;3,19,2024,groceries,w,54.82;5,03,2024,clothes,w,75.45;7,25,2024,entertainment,w,125.50', 3243))
>>> get_hash(stmt[0]) == '31967d1cf0bfa31a78960075d65ea84d' and len(stmt[1]) == 0 and (get_hash(stmt[2]) == get_hash(stmt[4]) == get_hash(stmt[-1]) == 'd856d631339828ebc2bde0f037fdaafe') and (get_hash(stmt[-2]) == '26cc06bd62dacaa34f0a37297f9fcc37')
True
```
```
>>> stmt = format_statement('ONWRBOIWER', build_transactions_list('2,9,2022,ATM Withdrawal,w,50;2,13,2024,paycheck,d,540.45;3,19,2024,groceries,w,54.82;5,03,2024,clothes,w,75.45;7,25,2024,entertainment,w,125.50', 3243))
>>> get_hash(stmt[0]) == '31967d1cf0bfa31a78960075d65ea84d' and get_hash(stmt[-2]) == '26cc06bd62dacaa34f0a37297f9fcc37'
True
```
