```python
# Initialize Otter
import otter
grader = otter.Notebook("lab10A.ipynb")
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session #10-A <br></center></h2>

<h1><center>Matrix operations<br></center></h1>

---

```python
from resources.hashutils import *
import numpy as np
import matplotlib.pyplot as plt
```

# Question 1: Matrix multiplication with Python

Use NumPy's `@` operator (which is a shorthand for the [matmul](https://numpy.org/doc/stable/reference/generated/numpy.matmul.html) function) to find the product $AB$ of the following two matrices:

$$A = \begin{bmatrix}0&1\\2&-1\end{bmatrix}$$
$$B = \begin{bmatrix}4&0&-1\\1&3&1\end{bmatrix}$$

Store the result as `Z`.

```python
A = ...
B = ...
Z = ...
```

```python
grader.check("q1p1")
```

# Question 2: Isometric projection

In this part we will compute the *isometric projection* that was covered in lecture. To recap, the isometric projection can be broken down into three steps. We begin with a coordinate frame $(x,y,z)$, and with the viewer located on the $x$ axis, and looking toward the origin. The three steps are then:
1. Rotate the scene by $\beta=45\degree$ on the $z$ axis. This produces a new coordinate frame $(x',y',z')$.
2. Rotate the scene by $\alpha=\text{atan}(1/\sqrt{2})$ on the $y'$ axis. This produces a new coordinate frame $(x'',y'',z'')$.
3. Relabel the coordinates: $x'''=y''$ and $y'''=z''$. Discard the $x''$ coordinate. 

All three of these steps are linear transformations, and can therefore be computed with matrix multiplications. In what follows we will execute the three steps in order. 

## Question 2.1: Load the point cloud of 3D points

The data consists of a cloud of points in space taken from the surface of an object. The data is stored in the file `pointcloud.npy`. Use NumPy's [load](https://numpy.org/doc/stable/reference/generated/numpy.load.html) method to load the data into a variable named `pointcloud`. 

Inspect the array to obtain the number of points in the cloud. Store this quantity as `numpoints`.

**Note**
+ The three rows in `pointcloud` are respectively the $x$, $y$, and $z$ coordinates of each point.

```python
pointcloud = ...
numpoints = ...
```

```python
grader.check("q2p1")
```

## Question 2.2: Plot the point cloud in 3D

Create a function called `plot_pc_3d` that takes a point cloud array as input (e.g. `pointcloud`) and produces a 3D scatter plot. The plot should have these characteristics:
+ `figsize=(6,6)`
+ For the scatter plot use a `'.'` of size 5 as the marker. 
+ Add a size 40 magenta-colored `'o'` marker at the origin (again using `scatter`)
+ Label the $x$, $y$, and $z$ axes with `'x'`, `'y'`, and, `'z'` respectively. 
+ The last line in the template (`ax.set_aspect('equal')`) gives equal aspect ratios to all three coordinates. This gives the scatter plot a realistic look. 

**Notes**:
+ The three rows in `pointcloud` are respectively the $x$, $y$, and $z$ coordinates of each point.
+ The function should return the `fig` and `ax` objects.
+ Applied to `pointcloud`, the plot should look like this:

<img src="resources/airplane1.png" width="400" />

+ Try adding `%matplotlib widget` to the cell that calls `plot_pc_3d(pointcloud)`. This will render the plot in a widget where you can rotate it with your mouse. But don't forget to remove this line before submitting, or else Gradescope's autograder will fail. 
 

```python
def plot_pc_3d(pts):
    fig = plt.figure(...)
    ax = fig.add_subplot(...)
    ...
    ... # etc...
    ax.set_aspect('equal')
    return fig, ax
```

```python
plot_pc_3d(pointcloud)
```

```python
grader.check("q2p2")
```

## Question 2.3 : Initial view

As mentioned earlier, the viewer is initially located on the positive $x$ axis, looking toward the origin. Their perspective is obtained by projecting the points onto the $yz$ plane. Create a function called `viewer_perspective` that plots this persepective. That is, the $y$ coordinate of each point should be plotted along the horizontal axis, and the $z$ coordinate on the vertical axis. Here are the specifications for the plot:

+ This is a 2D plot
+ `figsize=(6,6)`
+ Same marker as in Question 2.2
+ Again add a margent-colored `'o'` at the origin. 
+ Label the horizontal axis with `'y'` and the vertical axis with `'z'`.
+ Again use `ax.set_aspect('equal')`.
+ Return `fig, ax`
+ The plot should look like this:

<img src="resources/airplane2.png" width="400" />


```python
def viewer_perspective(pts):
    fig, ax = plt.subplots(figsize=(6,6))
    ...
    ...
    ...
    ax.set_aspect('equal')
    return fig, ax
```

```python
viewer_perspective(pointcloud)
```

```python
grader.check("q2p3")
```

## Question 2.4: $45\degree$ rotation about $z$

Step one in the process is to rotate the view $45\degree$ about the $z$ axis. This is a linear transformation from the $(x,y,z)$ coordinate frame to $(x',y',z')$, and it can be computed by multiplying each point by the matrix:
$$R1 = \begin{bmatrix} 
1/\sqrt{2} & 1/\sqrt{2}  & 0 \\ 
-1/\sqrt{2} & 1/\sqrt{2}  & 0 \\ 
0 & 0 & 1 \end{bmatrix} $$

Write a function called `doR1` that takes a `(3,numpoints)` shaped array as input, applies the $45\degree$ rotation to each point (by multiplying each one by $R1$), and returns the transformed point cloud (also shape `(3,numpoints)`). 

Use this function to compute a rotated point cloud:
```python 
pc1 = doR1(pointcloud)
```
And use `viewer_perspective(pc1)` to visualize the perspective of a viewer located on the rotated $x'$ axis. 

```python
def doR1(pts):
    R1 = ...
    ...
    return ...
```

```python
pc1 = doR1(pointcloud)
viewer_perspective(pc1)
```

```python
grader.check("q2p4")
```

## Question 2.5: Rotation about $y'$

The second step is to rotate the $(x',y',z')$ coordinate frame by $\alpha=\text{atan}(1/\sqrt{2})$ radians about the $y'$ axis. This is accomplished by multiplying each point in `pc1` by
$$R2 = \begin{bmatrix} 
\cos(\alpha) & 0 & \sin(\alpha) \\ 
0 & 1 & 0 \\ 
-\sin(\alpha) & 0 & \cos(\alpha) \end{bmatrix} $$

Write a function called `doR2` that transforms a point cloud from $(x',y',z')$ coordinates to $(x'',y'',z'')$ coordinates by multiplying each point by $R2$. The input and return values should be arrays of shape `(3,numpoint)`.

Then apply this transformation to `pc1` to obtain `pc2`, the point cloud after two rotations. Plotting `pc2` with  `viewer_perspective` produces the so-called "isometric projection" shown below.

<img src="resources/airplane3.png" width="400" />


```python
def doR2(pts):
    alpha = ...
    ...
    R2 = ...
    ...
    return ...
```

```python
pc2 = doR2(pc1)
viewer_perspective(pc2)
```

```python
grader.check("q2p5")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1p1 (points: None)
```
>>> get_hash(A, 1) == '18bfa568d8a267d82bada6a2740bcc8c' and get_hash(B, 1) == 'b9953b3a9d379e8b60ae8a6c18925c3b'
True
```
```
>>> get_hash(Z, 1) == '93293a2ce2ba3acb512c5590d2f4ec1e'
True
```


Tests for q2p1 (points: None)
```
>>> get_hash(pointcloud, 4) == '7dce19077e963e1e37aa86de220d2b3d'
True
```
```
>>> get_hash(numpoints, 4) == '0b8045e59d8cd9d75a263f5bcba540e4'
True
```


Tests for q2p2 (points: None)
```
>>> a = plt.ioff()
>>> import mpl_toolkits
>>> (fig, ax) = plot_pc_3d(np.array([[1, 2, 3], [3, 2, 1], [4, 5, 6]]))
>>> a = plt.ion()
>>> plt.close(fig)
>>> isinstance(ax, mpl_toolkits.mplot3d.axes3d.Axes3D)
True
```
```
>>> a = plt.ioff()
>>> (fig, ax) = plot_pc_3d(np.array([[1, 2, 3], [3, 2, 1], [4, 5, 6]]))
>>> a = plt.ion()
>>> plt.close(fig)
>>> np.all(np.sort(np.unique(np.vstack([ax.collections[i].get_offsets().data for i in range(2)]))) == np.array([0.0, 1.0, 2.0, 3.0]))
True
```
```
>>> a = plt.ioff()
>>> (fig, ax) = plot_pc_3d(np.array([[1, 2, 3], [3, 2, 1], [4, 5, 6]]))
>>> a = plt.ion()
>>> plt.close(fig)
>>> ax.get_xlabel() == 'x' and ax.get_ylabel() == 'y' and (ax.get_zlabel() == 'z')
True
```


Tests for q2p3 (points: None)
```
>>> a = plt.ioff()
>>> (fig1, ax1) = viewer_perspective(np.array([[1, 2, 3], [3, 2, 1], [4, 5, 6]]))
>>> a = plt.ion()
>>> assert np.all(np.sort(np.unique(np.vstack([ax1.collections[i].get_offsets().data for i in range(2)]))) == np.array([0.0, 1.0, 2.0, 3.0, 4.0, 5.0, 6.0]))
>>> plt.close(fig1)
```


Tests for q2p4 (points: None)
```
>>> pc1 = doR1(pointcloud)
>>> pc1.shape == (3, 2048)
True
```
```
>>> pc1 = doR1(pointcloud)
>>> get_hash(pc1, 4) == '96be8ec48ad0d58dc512726fedd7a938'
True
```


Tests for q2p5 (points: None)
```
>>> pc2 = doR2(doR1(pointcloud))
>>> pc2.shape == (3, 2048)
True
```
```
>>> pc2 = doR2(doR1(pointcloud))
>>> get_hash(pc2, 4) == 'aee700d10c6e58b2a82e0ff4101c0f3b'
True
```
